Elevated glycolytic metabolism of monocytes limits the generation of HIF-1α-driven migratory dendritic cells in tuberculosis

Author Response

The following is the authors’ response to the original reviews.

Reviewer #1 (Recommendations For The Authors):

Firstly, the authors place a great deal of emphasis on the impact of the Hif1-a inhibitor PX-478. The literature surrounding this inhibitor and its mode of action indicates that it is not a direct inhibitor of activity but that its greatest impact is on the production of Hif1-a. The authors do include another inhibitor as a control, Echinomycin, but it does not appear to be as biologically active and the panel of experiments conducted with this is extremely limited. I would be more comfortable with a full Seahorse experimental panel for Echinomycin, similar to SFig 2.G as performed with PX-478.

We thank the reviewer for their comment highlighting the different mechanisms of action of the HIF-1α inhibitors used in this article. While echinomycin inhibits the binding of HIF-1α to the hypoxia response element (HRE) thereby blocking HIF-1a DNA binding capability, PX-478 inhibits HIF-1α deubiquitination, decreases HIF-1α mRNA expression, and reduces HIF-1α translation. We have included a paragraph explaining this phenomenon in the new version of the manuscript (page 9). In addition, we extended the panel of experiments performed with echinomycin, which confirmed a marked inhibition of the glycolytic pathway when DCs were stimulated with irradiated Mtb in the presence of echinomycin as assessed by SCENITH (new Figure S3H).

Similarly, it would be of value to have Seahorse profiling that directly excludes FAO from the metabolic profile through the use of Etomoxir as an inhibitor of fatty acid oxidation, which one would assume would have no impact on the metabolic response.

In order to estimate the contribution of FAO towards fueling protein synthesis in DCs stimulated with iMtb, the FAO inhibitor etomoxir was incorporated to the SCENITH method as previously described (Adamik et al., 2022). Overall, FAO dependence was found to be less than 10% in DCs, regardless of their activation state. While mitochondrial dependence is reduced after iMtb stimulation, there is no difference in FAO dependence, suggesting that OXPHOS is primarily driven by glucose in iMtb-stimulated cells. This is consistent with HIF1α-induced increase of glucose metabolism-related genes. We have adjusted the results section to include this new result (new Figure S1).

Aside from these minor points, I believe this to be a rigorous study.

Reviewer #2 (Recommendations For The Authors):

In Fig. 1 and Fig. 2, the authors conclude that Mtb rewires the metabolism of Mo-DCs and induces both glycolysis and OXPHOS. The data shows that infection with iMtb or Mtb increases glucose uptake and lactate release, suggesting an increase in glycolysis. However, an increase in lactate is not a measure of glycolysis. Lactate is a byproduct of glycolysis; the end product of glycolysis is pyruvate.

We are grateful for the reviewer's comment, as it gives us the opportunity to explain the conceptual framework on which we based our study. Traditionally, pyruvate has been considered to be the end product of glycolysis when oxygen is present and lactate the end product under hypoxic conditions. Numerous studies have shown that lactate is produced even under aerobic conditions (Brooks, 2018). Therefore, we frame this work in accordance with this view that states that glycolysis begins with glucose as its substrate and terminates with the production of lactate as its main end product (Rogatzki, Ferguson, Goodwin, & Gladden, 2015; Schurr, 2023; Schurr & Schurr, 2017).

Secondly, since the authors have access to the Agilent Extracellular Flux Analyzer, they should have performed detailed ECAR/OCR measurements to conclusively demonstrate that both glycolysis and OXPHOS are increased in Mo. This is especially important for OXPHOS because the only readout shown for OXPHOS is an increase in mitochondrial mass (figure 1 G, H), which is not acceptable. Overall, the data does not indicate that Mtb triggers OXPHOS in the dendritic cells. It only indicates dead iMtb increases the mass of mitochondria in DCs.

The reviewer’s advice is well appreciated. However, we would like to clarify what may be a misunderstanding; that is, the assays alluded to by the reviewer were not performed on monocytes but on DCs. As advised by the reviewer, we now include the OCR measurements by Seahorse and describe the figures according to their order of appearance in the new version of the manuscript.

What happens to the mitochondrial mass when infected with live Mtb?

In response to the reviewer’s question, we determined the mitochondrial mass in infected DCs with live Mtb. In contrast to DCs treated with irradiated Mtb, those infected with live bacteria showed a clear reduction of their mitochondrial mass (modified Figure 1G). This result indicates that, although both Mtb-infected and irradiated Mtb-exposed DCs show a clear increase in their glycolytic activity, divergent responses are observed in terms of mitochondrial mass.

It will be best if the authors indicate in the figure headings that dead Mtb was used.

We agree with the reviewer. For figures 1-3, we applied the term “Mtb” in the figure headings since both irradiated and viable bacteria were used for the corresponding experiments. In figures 4-5, the term “iMtb” (alluding to irradiated Mtb) was used in the figure headings as suggested by the reviewer. For the remaining figures, the term “iMtb” was indicated in their legends when dead bacteria weres used to stimulate DCs.

E.g., Figure 1F; what does live Mtb do to GLUT1 levels etc etc?

In response to the reviewer’s question, we included new data about Glut1 expression in DCs infected with live Mtb in the latest version of the manuscript. In line with the increase in glucose uptake shown in figure 1B, we observed an increase in the percentage of Glut1 positive DCs upon Mtb infection (new Figure 1F, lower panels). The increase in Glut1 expression strengthens the notion that DCs activates their glycolytic activity in response to the infection, as demonstrated by the elevated release of lactate, glucose consumption, HIF-1α expression, LDHA expression (Figure 1) and glycolytic activity (Figure 2, SCENITH results with viable Mtb). Therefore, these data strongly support the induction of glycolysis by Mtb (either viable or irradiated) in DCs.

Also, we found that they were still able to activate CD4+ T cells from PPD+ donors in response to iMtb. This activation of CD4 T cells with iMtb in the presence of a HIF-1alpha inhibitor is expected, as iMtb is dead and not virulent. What happens when the cells are infected with live virulent Mtb?

We would like to clarify the main purpose of the DC-T cells co-culture assays in the presence of the HIF-1α inhibitors. To characterize the impact of HIF-1α on DC functionality, we assessed the capacity of DCs to activate autologous CD4+ T cells when stimulated with iMtb in the presence of HIF-1α inhibitors. To this end, we used iMtb merely as a source of antigens to load DCs and evaluate the effect of HIF-1α inhibition on the activation of antigen-specific T cell. The use of viable Mtb may introduce confounding factors, such as pathogen-triggered inhibitory mechanisms (e.g., EsxH secretion by Mtb, (Portal-Celhay et al., 2016)), which would prevent us from reaching conclusions about the role of HIF-1α. Thus, we consider that the use of live bacteria for this experiment is out of the scope of this manuscript.

The authors demonstrated that CD16+ monocytes from TB patients have higher glycolytic capacity than healthy controls Fig 7. The authors should differentiate TB patient monocytes into DCs and measure their bioenergetics to test if infection alters their glycolysis and OXPHOS.

In agreement with the reviewer, the determination of metabolic pathways in DCs differentiated from monocytes of TB patients is a key aspect of this work. Accordingly, the bioenergetic determinations of DCs generated from monocytes from TB patients versus healthy subjects are now illustrated in Figures 6F (lactate release) and 6G (SCENITH profile).

In the discussion, the authors state that "pathologically active glycolysis in monocytes from TB patients leads to poor glycolytic induction and migratory capacities of monocyte-derived DCs." However, the data from Fig. 1 and 2 show that treatment with iMtb or Mtb induces glycolysis in MoDCs. How do the authors explain these contrasting results?

We thank the reviewer for pointing out this issue. Figures 1 and 2 show DCs differentiated from monocytes of healthy donors (HS). In this case, DCs from HS respond to Mtb by inducing a glycolytic and migratory profile. Yet, in the case of monocytes isolated from TB patients, these cells exhibit an early glycolytic profile from the beginning of differentiation, ultimately yielding DCs with low glycolytic capacity and low migratory activity in response to Mtb. We included this explanation in the discussion (page 18) to better clarify this issue.

Also, the term "pathological" active glycolysis (Introduction and Discussion) is an inappropriate term.

As requested by the reviewer, we excluded the term “pathological” to describe the phenomenon reported in this study.

Lastly, it should be shown whether the DCs generated from CD16+ monocyte from TB patients generate tolerogenic and/or aberrant DCs, which have lower glycolytic and migration capacity compared to the CD16- monocyte population. In Figure 7B, the authors should discuss why the CD16+ monocyte population has lower glycolytic capacity compared to CD16- monocytes in healthy donors. Furthermore, in contrast to the TB patients, do DCs generated from CD16+ monocyte in healthy donors have increased glycolytic and migration capacity compared to CD16- monocyte (because these monocytes showed lower glycolytic capacity)? Furthermore, if there is no difference in glycolytic capacity among the three monocyte populations in TB patients, on what basis was it concluded that DCs generated only from the CD16+ monocyte population may be the cause of lower migration capacity? The authors state in Figure 7F that the DMOG pretreatment matches the situation where the Mo-DCs from TB patients showed reduced migration. Did the authors check the Hif-1alpha levels in monocytes obtained from TB patients?

We appreciate this in-depth analysis by the reviewer because it allows us to clarify some interpretations of the SCENITH results in Figure 7B. It is important to keep in mind that with the SCENITH technique we can only infer about the relative contributions between the metabolic pathways, without alluding to the absolute magnitudes of such contributions. In this regard, it is key to note that the amount of lactate released during the first hours of the TB monocyte culture is much higher than that released by monocytes from healthy subjects (HS, Figure 7A), even when most of monocytes, which are CD14+ CD16-, have comparable glycolytic capacities between HS and TB. Another example to illustrate how to interpret SCENITH results can be found in Figure 2, where a lower mitochondrial dependence is observed in iMtb-stimulated DCs (Figure 2A), while the absolute ATP production associated to OXPHOS is indeed higher as measured by Seahorse (Figure 2D). Therefore, the glycolytic capacity is not a direct readout of the magnitude of glycolysis, but of its contribution to total metabolism. The low levels of lactate released from HS monocytes likely reflects their low activation state and low metabolic activity compared to TB monocytes. In this regard, we have previously demonstrated that monocytes from pulmonary TB patients display an activated phenotype (Balboa et al., 2011). The fact that there is no difference between the glycolytic capacities of TB and HS CD16- monocytes indicates that their proportional contributions to protein synthesis are comparable (again, without inferring about their absolute values, which may be very different).

Beyond the previous clarification, the reviewer's proposal to isolate subsets of monocytes is a very interesting idea. However, the experimental approach is very difficult based on the amount of blood we can obtain from patients. The cohort of patients included in this work comprises very severe patients and we are given up to 15-20 ml of peripheral blood from each. This volume of blood yields up to 10 million PBMC with approximately 1 million monocytes. If we separate the monocyte subsets, the recovered cells per condition will be insufficient to perform the intended assays.

Nevertheless, we incorporate new evidence that TB disease is associated with an increased activation and glycolytic profile of circulating CD16+ monocytes.

i) First, we show that the baseline glycolytic capacity of CD16+ monocytes correlates with time since the onset of TB-related symptoms (new Figure 7C).

ii) Second, we performed high-throughput GeneSet Enrichment Analysis (GSEA) on transcriptomic data (GEO accession number: GSE185372) of CD14+CD16-, CD14+CD16+ and CD14dimCD16+ monocytes isolated from individuals with active TB, latent TB (IGRA+), as well as from TB negative healthy controls (IGRA-). We found enrichments that, unlike oxidative phosphorylation, glycolysis tends to increase in active TB in both CD14+CD16+ and CD14dimCD16+ monocytes (new Figure 7D).

iii) We measured the expression of HIF-1α in monocyte subsets by FACS and found that this transcription factor is expressed at higher levels in CD16+ monocyte subsets from TB patients compared to their counterparts from healthy donors (new Figure 8 A). We consider this result justifies the assays shown in Figure 8B-C, in which we prematurely activated HIF-1α in healthy donor monocytes during early differentiation to DCs and measured its impact on the migration of the generated DCs.

In the Discussion, the authors mention that circulating monocytes from TB patients differentiate from DCs with low immunogenic potential. However, the authors have not shown any immunological defect in any of their data with monocytes from TB patients. In the proxy model mentioned in Figure 7, they have in fact shown that these preconditioned DCs have higher CD86 expression. Can the authors explain/show data to justify the statement in the first paragraph of the Discussion?

We agree with the reviewer on this observation. Our findings are limited to the generation of DCs with low migratory potential (low chemotactic activity towards CCL21 of DC differentiated from TB patient monocytes shown in figure 6H and of DC generated from pre-conditioned monocytes shown in figure 8C). We have modified that part of the discussion to better clarify this point, replacing migratory with immunogenic.

The authors should note that oxamate is a competitive inhibitor of the enzyme lactate dehydrogenase and not glycolysis. Also, LDHA catalyzes the conversion from pyruvate to lactate and not the other way around (Results, page 6).

This comment relates to the first one by the reviewer, in which the dogma of glycolysis was discussed. According to the new conception of glycolysis, it begins with glucose as its substrate and terminates with the production of lactate as its main end product.

The following statements by the authors on page 6 are incorrect: "Because irradiated and viable Mtb induced comparable activation of glycolysis, we subsequently performed all our assays with irradiated Mtb only in the rest of the study due to biosafety reasons." and: "To our knowledge, this is the first study addressing the metabolic status and migratory activity of Mo-DCs from TB patients."

We deleted the first sentence and reworded the second sentence as "To our knowledge, this is the first study to address how the metabolic status of monocytes from TB patients influences the migratory activity of further differentiated DCs".

The Discussion reads as if live Mtb was used in the experiments, which is not the case. This should be corrected.

We changed Mtb for iMtb when it was the case in the discussion. In some cases, Mtb stimulation was used instead of Mtb infection.

Minor Comments:

(1) In Figure 1F legend "Quantification of Glut1+ cells plotted to the right". The underlined part should be "plotted below".

It was corrected.

(2) In Figure 1H. Please describe the quantitation method and describe how many cells or the number/size of fields were used to quantitate mitochondria.

For mitochondrial morphometric analysis, TEM images were quantified with the ImageJ “analyze particles” plugin in thresholded images, with size (μm2) settings from 0.001 to infinite. For quantification, 8–10 cells of random fields (1000x magnification) per condition were analyzed. We included this information in the methods section of the new version of the manuscript.

(3) Please mention the number of independent experimental repeats for each experimental data set and figure.

In each figure, the number of independent experiments is indicated by individual dots.

(4) In Figure 2A legend, "PER; left panel" should be PER; lower panel and "OCR; right panel" should be OCR; upper panel.

It was corrected.

References for reviewers

Adamik, J., Munson, P. V., Hartmann, F. J., Combes, A. J., Pierre, P., Krummel, M. F., … Butterfield, L. H. (2022). Distinct metabolic states guide maturation of inflammatory and tolerogenic dendritic cells. Nature Communications 2022 13:1, 13(1), 1–19. https://doi.org/10.1038/s41467-022-32849-1

Balboa, L., Romero, M. M., Basile, J. I., Sabio y Garcia, C. A., Schierloh, P., Yokobori, N., … Aleman, M. (2011). Paradoxical role of CD16+CCR2+CCR5+ monocytes in tuberculosis: efficient APC in pleural effusion but also mark disease severity in blood. Journal of Leukocyte Biology. https://doi.org/10.1189/jlb.1010577

Brooks, G. A. (2018). Cell Metabolism The Science and Translation of Lactate Shuttle Theory. Cell Metab. https://doi.org/10.1016/j.cmet.2018.03.008

Portal-Celhay, C., Tufariello, J. M., Srivastava, S., Zahra, A., Klevorn, T., Grace, P. S., … Philips, J. A. (2016). Mycobacterium tuberculosis EsxH inhibits ESCRT-dependent CD4+ T-cell activation. Nature Microbiology, 2, 16232. https://doi.org/10.1038/NMICROBIOL.2016.232

Rogatzki, M. J., Ferguson, B. S., Goodwin, M. L., & Gladden, L. B. (2015). Lactate is always the end product of glycolysis. Frontiers in Neuroscience, 9(FEB), 125097. https://doi.org/10.3389/FNINS.2015.00022/BIBTEX

Schurr, A. (2023). From rags to riches: Lactate ascension as a pivotal metabolite in neuroenergetics. Frontiers in Neuroscience, 17, 1145358. https://doi.org/10.3389/FNINS.2023.1145358/BIBTEX

Schurr, A., & Schurr, A. (2017). Lactate, Not Pyruvate, Is the End Product of Glucose Metabolism via Glycolysis. Carbohydrate. https://doi.org/10.5772/66699

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: During tuberculosis, migration of dendritic cells (DCs) from the site of infection to the draining lymph nodes is known to be impaired, hindering the rapid development of protective T-cell mediated immunity. However, the mechanisms involved in the delayed migration of DCs during TB are still poorly defined. Here, we found that infection of DCs with Mycobacterium tuberculosis triggers HIF-1α-mediated aerobic glycolysis in a TLR2-dependent manner, and that this metabolic profile is essential for DC migration. In particular, oxamate, a glycolysis inhibitor, or PX-478, an HIF-1α inhibitor, completely abrogated M. tuberculosis-induced DC migration in vitro to the lymphoid tissue chemokine CCL21, and in vivo to lymph nodes in mice. Strikingly, we found that although monocytes from TB patients are inherently biased toward glycolysis metabolism, they differentiate into poorly glycolytic and poorly migratory DCs, compared with healthy subjects. Taken together, these data suggest that because of their preexisting glycolytic state, circulating monocytes from TB patients are refractory to differentiation into migratory DCs, which may explain the delayed migration of these cells during the course of the disease and opens avenues for host-directed therapies for TB.

Revised version: During tuberculosis, migration of dendritic cells (DCs) from the site of infection to the draining lymph nodes is known to be impaired, hindering the rapid development of protective T-cell mediated immunity. However, the mechanisms involved in the delayed migration of DCs during TB are still poorly defined. Here, we found that infection of DCs with Mycobacterium tuberculosis triggers HIF-1α-mediated aerobic glycolysis in a TLR2-dependent manner, and that this metabolic profile is essential for DC migration. In particular, the glycolysis inhibitor oxamate and the HIF-1α inhibitor PX-478 abrogated M. tuberculosis-induced DC migration in vitro to the lymphoid tissue-specific chemokine CCL21, and in vivo to lymph nodes in mice. Strikingly, we found that although monocytes from TB patients are inherently biased toward glycolysis metabolism, they differentiate into poorly glycolytic and poorly migratory DCs, compared with healthy subjects. Taken together, these data suggest that because of their preexisting glycolytic state, circulating monocytes from TB patients are refractory to differentiation into migratory DCs, which may explain the delayed migration of these cells during the disease and opens avenues for host-directed therapies for TB.

Introduction

Earlier version: Tuberculosis (TB) remains a major global health problem, responsible for approximately 1.5 million deaths annually. The causative agent of TB, Mycobacterium tuberculosis (Mtb), is a highly successful pathogen that has evolved several strategies to weaken the host immune response. Although reliable immune correlates of protective immunity against Mtb are still not well-defined, it is widely accepted that Th1 cells contribute to protection by secreting IFN-γ and promoting antimycobacterial activity in macrophages1. Importantly, the induction of a strong Th1 immune response relies on the generation of immunogenic dendritic cells (DCs) with strong migratory properties2–5. Mtb has been shown to interfere with several DC functions, thus impairing the induction and development of adaptive immunity6–8. For instance, we and others previously reported that Mtb-exposed DCs have low capacity for mycobacterial antigen presentation and stimulation of Mtb-specific CD4+ T cells9–13. Additionally, Mtb-infected DCs were reported to have impaired ability to migrate to lymph nodes in vitro14,15 and in murine models3,5; however, the underlying molecular mechanisms of these phenotypes and their relevance to the migratory activity of monocyte-derived DCs in TB patients remain unknown.

Revised version: Tuberculosis (TB) remains a major global health problem, responsible for approximately 1.6 million deaths annually. The causative agent of TB, Mycobacterium tuberculosis (Mtb), is a highly successful pathogen that has evolved several strategies to weaken the host immune response. Although reliable immune correlates of protective immunity against Mtb are still not well-defined, it is widely accepted that Th1 cells contribute to protection by secreting IFN-γ and promoting antimycobacterial activity in macrophages1. Importantly, the induction of a strong Th1 immune response relies on the generation of immunogenic dendritic cells (DCs) with strong migratory properties2–5. Mtb has been shown to interfere with several DC functions, thus impairing the induction and development of adaptive immunity6–8. For instance, we and others previously reported that Mtb-exposed DCs have low capacity for mycobacterial antigen presentation and stimulation of Mtb-specific CD4+ T cells9–13. Additionally, Mtb-infected DCs were reported to have an impaired ability to migrate to lymph nodes in vitro14, 15 and in vivo in murine models3, 5; however, the underlying molecular mechanisms of these phenotypes and their relevance to the migratory activity of monocyte-derived DCs in TB patients remain unknown.

Earlier version: Rapid, directed migration of DCs towards secondary lymphoid organs requires essential changes at the cellular and molecular levels16. Relatedly, the metabolic state of DCs is complex and varies according to cell origin, differentiation and maturation state as well as local microenvironment, among other factors17–20. Studies have reported that upon pathogen sensing, the transcription factor hypoxia-inducible factor-1α (HIF-1α) increases glycolysis, which promotes immunogenic functions of DCs, such as IL-12 production, costimulatory marker expression21, and cell migration22–24. By contrast, it was shown that HIF-1α represses the proinflammatory output of LPS-stimulated DCs and can inhibit DC-induced T-cell responses in other settings25. To reconcile these disparate roles for HIF-1α, it has been proposed that the impact of metabolic pathway activation on DC function depends on DC subset18. To this point, most prior studies have been conducted using murine conventional DCs and plasmacytoid DCs19. Recently, with the implementation of high-dimensional techniques, it was demonstrated that distinct metabolic wiring is associated with individual differentiation and maturation stages of DCs26, highlighting the importance of defining the metabolic profile of specific subsets of DCs under particular physiological or pathological conditions20. Given the key role of DCs in the host response to TB, it is thus crucial to investigate DC metabolism in the context of Mtb infection27.

Revised version: Rapid, directed migration of DCs towards secondary lymphoid organs requires essential changes at the cellular and molecular levels16. Relatedly, the metabolic state of DCs is complex and varies according to cell origin, differentiation and maturation states, as well as local microenvironment, among other factors17–20. Studies have reported that upon pathogen sensing, the transcription factor hypoxia-inducible factor-1α (HIF-1α) increases glycolysis, which promotes immunogenic functions of DCs, such as IL-12 production, costimulatory marker expression21, and cell migration22–24. By contrast, it was shown that HIF-1α represses the proinflammatory output of LPS-stimulated DCs and can inhibit DC-induced T-cell responses in other settings25. To reconcile these disparate roles for HIF-1α, it has been proposed that the impact of metabolic pathway activation on DC functions varies among DC subsets18. To this point, most prior studies have been conducted using murine conventional DCs and plasmacytoid DCs19. Recently, with the implementation of high-dimensional techniques, it was demonstrated that distinct metabolic wiring is associated with individual differentiation and maturation stages of DCs26, highlighting the importance of defining the metabolic profile of specific subsets of DCs under physiological or pathological conditions20. Given the key role of DCs in the host response to TB, it is thus crucial to investigate DC metabolism in the context of Mtb infection27.

Earlier version: We previously demonstrated that the TB-associated microenvironment, as conferred by the acellular fraction of TB patient pleural effusions, inhibits HIF-1α activity leading to a reduction in glycolytic and microbicidal phenotypes in macrophages28. Moreover, activation of HIF-1α enhances Mtb control at early times post-infection in mouse models29, and this effect was associated with a metabolic switch of alveolar macrophages towards an M1-like profile28. Given that HIF-1α activation promotes protection at early stages of Mtb infection and given its role as a key regulator of DC migration and inflammation30, we hypothesized that HIF-1α could affect the functionality of DCs in regulating the initiation and orchestration of the adaptive immune response to Mtb, a process known to be delayed upon Mtb infection5,6. Here, we show that HIF-1α-mediated glycolysis promotes DC activation and migration in the context of TB. Importantly, we report pathologically active glycolysis in monocytes from TB patients, which leads to poor glycolytic induction and migratory capacities of monocyte-derived DCs.

Revised version: We previously demonstrated that the TB-associated microenvironment, as conferred by the acellular fraction of TB patient pleural effusions, inhibits HIF-1α activity leading to a reduction in glycolytic and microbicidal phenotypes in macrophages28. Moreover, activation of HIF-1α enhances Mtb control at early times post-infection in mouse models29, and this effect was associated with a metabolic switch of alveolar macrophages towards an M1-like profile28. Given that HIF-1α activation promotes protection at early stages of Mtb infection and given its role as a key regulator of DC migration and inflammation30, we hypothesized that HIF-1α could affect the functionality of DCs in regulating the initiation and orchestration of the adaptive immune response to Mtb, a process known to be delayed upon Mtb infection5, 6. Here, we show that HIF-1α-mediated glycolysis promotes DC activation and migration in the context of TB. Importantly, we report active glycolysis in monocytes from TB patients, which leads to poor glycolytic induction and migratory capacities of monocyte-derived DCs.

Results

Mtb impacts metabolism in human monocyte-derived DCs

Earlier version: To determine the impact of Mtb on the metabolism of human monocyte-derived DCs (Mo-DCs), we assessed metabolic parameters associated either with glycolysis or oxidative phosphorylation (OXPHOS) upon Mtb stimulation. Cells undergoing aerobic glycolysis are characterized by increased consumption of glucose and the production and release of lactate. We measured lactate release and glucose consumption in Mo-DCs stimulated for 24h with equivalent doses of either irradiated (iMtb) or viable Mtb. DCs treated with either iMtb or viable Mtb released increased levels of lactate and consumed more glucose than untreated DCs (Figure 1A-B). Consistently, both iMtb treatment and Mtb infection resulted in an increase in expression of the key glycolysis-activating regulator HIF-1α at both mRNA and protein levels (Figure 1C-D). Expression of the gene encoding the glycolytic enzyme lactate dehydrogenase A (LDHA), which catalyzes the conversion of lactate to pyruvate, was also increased in iMtb-treated or Mtb-infected DCs (Figure 1E). In agreement with their enhanced glycolysis profile, DCs stimulated with iMtb had increased expression of the glucose transporter GLUT1 (SLC2A1)36 (Figure 1F). Of note, LDHA and GLUT1 are HIF-1α target genes, and their upregulation correlated with the increase in HIF-1α expression upon Mtb stimulation. Because irradiated and viable Mtb induced comparable activation of glycolysis, we subsequently performed all our assays with irradiated Mtb only in the rest of the study due to biosafety reasons. To assess changes in the mitochondria, we measured mitochondrial mass and morphology upon iMtb stimulation. We found a higher mitochondrial mass as well as larger individual mitochondria in iMtb-stimulated DCs compared to untreated DCs (Figure 1G-H). Thus, our data indicate that Mtb triggers both glycolysis and OXPHOS in Mo-DCs.

Revised version: To determine the impact of Mtb on the metabolism of human monocyte-derived DCs (Mo-DCs), we assessed metabolic parameters associated with glycolysis and mitochondrial changes upon Mtb stimulation or infection. Cells undergoing aerobic glycolysis are characterized by increased consumption of glucose and the production and release of lactate. We measured lactate release and glucose consumption in Mo-DCs stimulated for 24 h with equivalent doses of either irradiated (iMtb) or viable Mtb. DCs treated with either iMtb or viable Mtb released increased levels of lactate and consumed more glucose than untreated DCs (Figure 1A-B). Consistently, both iMtb treatment and Mtb infection resulted in an increase in expression of the key glycolysis-activating regulator HIF-1α at both mRNA and protein levels (Figure 1C-D). Expression of the gene encoding the glycolytic enzyme lactate dehydrogenase A (LDHA), which catalyzes the conversion of lactate to pyruvate, was also increased in iMtb-treated or Mtb-infected DCs (Figure 1E). In agreement with their enhanced glycolysis profile, DCs stimulated with iMtb or infected with viable Mtb had increased expression of the glucose transporter GLUT1 (SLC2A1)31 (Figure 1F). Of note, LDHA and GLUT1 are HIF-1α target genes, and their upregulation correlated with the increase in HIF-1α expression upon Mtb stimulation. To assess changes in the mitochondria, we measured mitochondrial mass and morphology. We found a higher mitochondrial mass as well as larger individual mitochondria in iMtb-stimulated DCs compared to untreated DCs (Figure 1G-H). In contrast to the findings obtained upon iMtb stimulation, Mtb-infected DCs displayed a reduction in their mitochondrial mass (Figure 1G). This result indicates that although both Mtb-infected and irradiated Mtb-exposed DCs show a clear increase in their glycolytic activity, divergent responses are observed in terms of mitochondrial mass. Therefore, our data indicate that Mtb impacts the metabolism of Mo-DCs, leading to mitochondrial changes and triggering glycolysis-associated parameters.

Mtb rewires the metabolic network of monocyte-derived DCs (Mo-DCs).

Earlier version: Mo-DCs were stimulated with viable or irradiated Mtb (iMtb) at two MOI (1 or 2 Mtb per DC) for 24h. Glycolysis was measured as: (A) Lactate release in culture supernatants; (B) Glucose uptake measured in culture supernatants; (C) Relative expression of HIF-1α mRNA normalized to EeF1A1 control gene. (D) Representative histograms of the mean fluorescence intensity (MFI) of HIF-1α as measured by flow cytometry. Quantification shown in graph to the right. (E) Relative expression of lactate dehydrogenase A (LDHA) mRNA normalized to EeF1A1 control gene. (F) FACS plots show the percentage of Glut1+ cells with and without iMtb stimulation in a representative experiment. Quantification of Glut1+ cells plotted to the right. (G) MFI of Mitospy probe as a measurement of mitochondrial mass for Mo-DCs with and without iMtb treatment. (H) Representative electron microscopy micrographs of control and iMtb-stimulated DCs showing mitochondria colored in cyan (left panels) and quantified morphometric analysis (right panels). Statistical significance was assessed in (A-E) using 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01; ∗∗∗∗p < 0.0001), and in (F-H) using paired t test (∗p < 0.05) for iMtb versus controls. All values are expressed as means ± SEM.

Revised version: Mo-DCs were stimulated with viable or irradiated Mtb (iMtb) at two MOI (1 or 2 Mtb per DC) for 24 h. Glycolysis was measured as: (A) Lactate release in culture supernatants; (B) Glucose uptake measured in culture supernatants; (C) Relative expression of HIF-1α mRNA normalized to EeF1A1 control gene. (D) Representative histograms of the mean fluorescence intensity (MFI) of HIF-1α as measured by flow cytometry. Quantification shown in graph to the right. (E) Relative expression of lactate dehydrogenase A (LDHA) mRNA normalized to EeF1A1 control gene. (F) FACS plots show the percentage of Glut1+ cells with and without iMtb stimulation or infected with viable Mtb in a representative experiment. Quantification of Glut1+ cells plotted below. (G) MFI of Mitospy probe as a measurement of mitochondrial mass for Mo-DCs treated (or not) with iMtb (upper panel) or infected with viable Mtb (lower panel). The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown. (H) Representative electron microscopy micrographs of control and iMtb-stimulated DCs showing mitochondria colored in cyan (left panels) and quantified morphometric analysis (right panels). Statistical significance was assessed in (A-E) using 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01; ∗∗∗∗p < 0.0001), and in (F-H) using paired t test (∗p < 0.05) for iMtb versus controls. All values are expressed as means ± SEM.

Earlier version: Mtb exposure shifts DCs to a glycolytic profile over OXPHOS

Revised version: Mtb exposure shifts DCs to a glycolytic profile over oxidative phosphorylation

Earlier version: To further characterize the metabolic profile of iMtb-stimulated DCs, we next assessed the intracellular rates of glycolytic vs. mitochondrial ATP production using Seahorse technology. Bioenergetic profiles revealed that iMtb increased the rate of protons extruded over time, or proton efflux rate (PER) as well as the basal oxygen consumption rate (OCR) in Mo-DCs (Figure 2A). The measurements of basal extracellular acidification rate (ECAR) and OCR were used to calculate ATP production rate from glycolysis (GlycoATP) and mitochondrial oxidative phosphorylation (MitoATP). In line with our earlier findings, the ATP production rates from both glycolysis and mitochondrial respiration were augmented upon iMtb-stimulation (Figure 2B). Of note, the relative contribution of GlycoATP to overall ATP production was increased, while MitoATP contribution was decreased in iMtb-treated cells compared to untreated cells (Figure 2C). To further check the heterogeneity, we also evaluated the metabolism of DCs at single cell level using the SCENITH technology35. This method is based on the fact that a decrease in ATP levels is tightly coupled with a decrease in protein synthesis and display similar kinetics35. By treating the cells with glucose or mitochondrial respiration inhibitors and measuring their impact on protein synthesis by puromycin incorporation via flow cytometry, glucose and mitochondrial dependences can be quantified. Two additional derived parameters, glycolytic capacity and fatty acids and amino acids oxidation (FAO & AAO) capacity were also calculated. Similarly to Seahorse results, SCENITH technology revealed a lower reliance on OXPHOS in parallel with an increase in the glycolytic capacity of Mtb-stimulated DCs (Figure 2D-E). No differences were observed for glucose dependence and FAO & AAO capacity (Figure 2D-E). These results confirmed the reconfiguration of DC metabolism by iMtb treatment, with an increase in the relative glycolytic contribution to overall metabolism at the expense of the OXPHOS pathway. Together, metabolic profiling indicates that a metabolic switch toward aerobic glycolysis occurs in Mo-DCs exposed to Mtb.

Revised version: To further characterize the metabolic profile of DCs upon iMtb-stimulation or Mtb-infection, we next evaluated the metabolism of DCs at single-cell level using the SCENITH technology32. This method is based on a decrease in ATP levels that is tightly coupled with a decrease in protein synthesis and displays similar kinetics32. By treating the cells with glucose or mitochondrial respiration inhibitors, and measuring their impact on protein synthesis by puromycin incorporation via flow cytometry, glucose and mitochondrial dependences can be quantified. Two additional derived parameters such as “glycolytic capacity” and “fatty acid and amino acid oxidation (FAO & AAO) capacity” were also calculated. SCENITH technology revealed a lower reliance on oxidative phosphorylation (OXPHOS) in parallel with an increase in the glycolytic capacity of iMtb-stimulated (Figure 2A), Mtb-infected DCs and even bystander DCs. (Figure 2-B). No differences were observed for glucose dependence and FAO & AAO capacity (Figure 2A-B). Additionally, we found no changes between the FAO dependency in Mtb-stimulated DCs in comparison to control cells when the FAO inhibitor (etomoxir) was used (Figure S1A). For the case of iMtb-stimulated DCs, we also assessed the intracellular rates of glycolytic and mitochondrial ATP production using Seahorse technology. Bioenergetic profiles revealed that iMtb increased the rate of protons extruded over time, or proton efflux rate (PER), as well as the basal oxygen consumption rate (OCR) in Mo-DCs (Figure 2C). The measurements of basal extracellular acidification rate (ECAR) and OCR were used to calculate ATP production rate from glycolysis (GlycoATP) and mitochondrial OXPHOS (MitoATP). The ATP production rates from both glycolysis and mitochondrial respiration were augmented upon iMtb-stimulation (Figure 2D). Similar to SCENITH results, the relative contribution of GlycoATP to overall ATP production was increased, while MitoATP contribution was decreased in iMtb-treated cells compared to untreated cells (Figure 2E). These results confirmed the change in DC metabolism induced by Mtb, with an increase in the relative glycolytic contribution to overall metabolism at the expense of the OXPHOS pathway. Together, metabolic profiling indicates that a metabolic switch toward aerobic glycolysis occurs in Mo-DCs exposed to Mtb.

Mtb skews DC metabolism toward glycolysis.

Earlier version: Mo-DCs were stimulated with irradiated Mtb (iMtb). (A) Kinetic profile of proton efflux rate (PER; left panel) and oxygen consumption rate (OCR; right panel) measurements in control and iMtb-stimulated DCs in response to inhibitor treatments (Oligomycin, O; ROT/AA, Rotenone/Antimycin A), obtained using an Agilent Seahorse XFe24 Analyzer. PER and OCR measurements were normalized to the area covered by cells. (B) ATP production rate from mitochondrial oxidative phosphorylation (MitoATP) and glycolysis (glycoATP). MitoATP production rate and glycoATP production rate were calculated from OCR and ECAR measurements in control and iMtb-stimulated DCs. (C) Percentages of MitoATP and GlycoATP relative to overall ATP production. (D) Representative histograms showing the translation level after puromycin (Puro) incorporation and staining with a monoclonal anti-Puro (anti-Puro MFI) in response to inhibitor treatment (C, Control; DG, 2-Deoxy-D-Glucose; Oligomycin, O; or combination treatment, DG+O). The bar plots show the values of the anti-Puro MFI from 6 donors. Arrows and numbers inside boxes denote the differences between the MFI of puro in the different treatments that are used to calculate the glucose dependence (1) and fatty acids and amino acids oxidation (FAO & AAO) capacity (4); and the mitochondrial dependency (2) and glycolytic capacity (3). (E) Relative contributions of glycolytic and FAO & AAO capacities and glucose and mitochondrial dependences to overall DC metabolism analyzed with SCENITH. Paired t test (∗p < 0.05; ∗∗p < 0.01) for iMtb versus controls. Values are expressed as means ± SEM.

Revised version: Mo-DCs were stimulated with irradiated Mtb (iMtb) or infected with Mtb expressing Red Fluorescent Protein (Mtb-RFP, panel C). (A) Representative histograms showing the translation level after puromycin (Puro) incorporation and staining with a monoclonal anti-Puro (anti-Puro MFI) in response to inhibitor treatment (C, Control; DG, 2-Deoxy-D-Glucose; Oligomycin, O; or combination treatment, DG+O). The bar plots show the values of the anti-Puro MFI from 6 donors. Arrows and numbers inside boxes denote the differences between the MFI of puro in the different treatments that are used to calculate the glucose dependence (1) and fatty acids and amino acids oxidation (FAO & AAO) capacity (4); and the mitochondrial dependency (2) and glycolytic capacity (3). (B) Relative contributions of glycolytic and FAO & AAO capacities and glucose and mitochondrial dependences to overall DC metabolism analyzed with SCENITH. (C) DCs were infected with Mtb-RFP for 24 h, thereafter the metabolic profile was evaluated by SCENITH. Representative plots showing the gating strategy to distinguish the populations within Mtb-infected cultures, which includes RFP+ (Mtb-infected DCs) and RFP- (bystander DCs) cells. Representative histograms showing the translation level after Puro incorporation are shown for uninfected, Mtb-infected and bystander DCs. The bar plots show the values of the anti-Puro MFI from 4 donors. Right panels show the relative contributions of glycolytic and FAO & AAO capacities and glucose and mitochondrial dependences to DC metabolism. (D) Kinetic profile of proton efflux rate (PER; lower panel) and oxygen consumption rate (OCR; upper panel) measurements in control and iMtb-stimulated DCs in response to inhibitor treatments (Oligomycin, O; ROT/AA, Rotenone/Antimycin A), obtained using an Agilent Seahorse XFe24 Analyzer. PER and OCR measurements were normalized to the area covered by cells. (E) ATP production rate from mitochondrial oxidative phosphorylation (MitoATP) and glycolysis (glycoATP). MitoATP production rate and glycoATP production rate were calculated from OCR and ECAR measurements in control and iMtb-stimulated DCs. (F) Percentages of MitoATP and GlycoATP relative to overall ATP production. Statistics in (B, E-F) are from paired t test (∗p < 0.05; ∗∗p < 0.01) for iMtb versus controls. Statistics in (C) are 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05) as depicted by lines. The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown.

Mtb triggers the glycolytic pathway through TLR2 ligation

Earlier version: Since Mtb is sensed by Toll-like receptors (TLR) 2 and 437, we investigated the contribution of these receptors to glycolysis activation in Mo-DCs upon Mtb stimulation. Using specific neutralizing antibodies for these receptors, we found that TLR2 ligation, but not that of TLR4, was required to trigger the glycolytic pathway, as reflected by a decrease in lactate release, glucose consumption and HIF-1α expression in iMtb-stimulated DCs treated with an anti-TLR2 antibody (Figure 3A-C). As a control, and as expected given the reliance on TLR4 for LPS sensing38, lactate release and glucose consumption were abolished in LPS-stimulated DCs in the presence of neutralizing antibodies against TLR4 but not TLR2 (Figure S1A-B). Moreover, blockade of TLR2 also diminished glycolytic ATP production in iMtb-stimulated DCs without altering OXPHOS-associated ATP production (Figure 3D) or the size and morphology of mitochondria (Figure S1C), suggesting that TLR2 engagement by iMtb is required for the induction of glycolysis but not mitochondrial respiration. Interestingly, TLR2 ligation was also necessary for lactate release triggered by viable Mtb, (Fig S1D). To further confirm the involvement of TLR2 in the induction of glycolysis, we tested the effect of synthetic (Pam3CSK4) or mycobacterial (peptidoglycans, PTG) TLR2 agonists39,40 and found that both ligands induced lactate release and glucose consumption in DCs (Figure 3E-F), without affecting cell viability (Figure S1E). Thus, our data indicate that Mtb induce glycolysis in Mo-DCs through TLR2 engagement.

Revised version: Since Mtb is sensed by Toll-like receptors (TLR)-2 and -433, we investigated the contribution of these receptors to glycolysis activation in Mo-DCs upon Mtb stimulation. Using specific neutralizing antibodies for these receptors, we found that TLR2 ligation, but not that of TLR4, was required to trigger the glycolytic pathway, as reflected by a decrease in lactate release, glucose consumption and HIF-1α expression in iMtb-stimulated DCs treated with an anti-TLR2 antibody (Figure 3A-C). As a control, and as expected given the reliance on TLR4 for LPS sensing34, lactate release and glucose consumption were abolished in LPS-stimulated DCs in the presence of neutralizing antibodies against TLR4 but not TLR2 (Figure S2A-B). Moreover, blockade of TLR2 also diminished glycolytic ATP production in iMtb-stimulated DCs without altering OXPHOS-associated ATP production (Figure 3D) or the size and morphology of mitochondria (Figure S2C), suggesting that TLR2 engagement by iMtb is required for the induction of glycolysis but not mitochondrial respiration. Interestingly, TLR2 ligation was also necessary for lactate release triggered by viable Mtb, (Figure S2D). To further confirm the involvement of TLR2 in the induction of glycolysis, we tested the effect of synthetic (Pam3CSK4) or mycobacterial (peptidoglycans, PTG) TLR2 agonists35, 36 and found that both ligands induced lactate release and glucose consumption in DCs (Figure 3E-F), without affecting cell viability (Figure S2E). Thus, our data indicate that Mtb induces glycolysis in Mo-DCs through TLR2 engagement.

Mtb triggers glycolysis through TLR2 ligation in Mo-DCs.

Earlier version: Mo-DCs were stimulated with irradiated Mtb (iMtb) in the presence of neutralizing antibodies against either TLR2 (aTLR2), TLR4 (aTLR4), or their respective isotype controls. (A) Lactate release as measured in supernatant. (B) Glucose uptake as measured in supernatant. (C) Mean fluorescence intensity (MFI) of HIF-1α as measured by flow cytometry. (D) Kinetic profile of proton efflux rate (PER) and oxygen consumption rate (OCR) measurements (left panels). Metabolic flux analysis showing quantification of mitochondrial ATP production and glycolytic ATP production (right panel). (E-F) Mo-DCs were stimulated with Pam3Cys or Mtb peptidoglycan (PTG) at the indicated concentrations. (E) Lactate release as measured in supernatant. (F) Glucose uptake as measured in supernatant. Statistics in (A-B, E-F) are 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01; ∗∗∗∗p < 0.0001). Statistics in (C-D) are from paired t test (∗p < 0.05) for iMtb versus controls. Values are expressed as means ± SEM.

Revised version: Mo-DCs were stimulated with irradiated Mtb (iMtb) in the presence of neutralizing antibodies against either TLR2 (aTLR2), TLR4 (aTLR4), or their respective isotype controls. (A) Lactate release as measured in supernatant. (B) Glucose uptake as measured in supernatant. (C) Mean fluorescence intensity (MFI) of HIF-1α as measured by flow cytometry. (D) Kinetic profile of proton efflux rate (PER) and oxygen consumption rate (OCR) measurements (left panels). Metabolic flux analysis showing quantification of mitochondrial ATP production and glycolytic ATP production (right panel). (E-F) Mo-DCs were stimulated with Pam3Cys or Mtb peptidoglycan (PTG) at the indicated concentrations. (E) Lactate release as measured in supernatant. (F) Glucose uptake as measured in supernatant. Statistics in (A-B, E-F) are 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01; ∗∗∗∗p < 0.0001). Statistics in (C-D) are from paired t test (∗p < 0.05) for iMtb versus controls. The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown.

HIF-1α is required for DC maturation upon iMtb stimulation but not for CD4+ T lymphocyte polarization

Earlier version: To determine the impact of glycolysis on DC maturation and the capacity to activate T cells, we inhibited HIF-1α activity in iMtb-stimulated DCs. Treatment with the HIF-1α inhibitor PX-478 (PX) diminished lactate release and glucose consumption in iMtb-stimulated DCs without affecting cell viability at the indicated concentration (Figure S2A-C). These results were confirmed using echinomycin (Ech), another HIF-1α inhibitor (Figure S2D-F). HIF-1α inhibition by PX significantly abolished ATP production associated with glycolysis without affecting absolute levels of OXPHOS-derived ATP production in iMtb-stimulated DCs (Figure 4A and Figure S2G). Although HIF-1α inhibitors did not affect the uptake of iMtb by DCs (Figure S2H), we observed a reduction in the expression of activation markers CD83 and CD86, but not in the inhibitory molecule PD-L1, upon treatment with PX (Figure 4B) or Ech (Figure S3). We then measured cytokine production in iMtb-stimulated DCs after HIF-1α inhibition and noted a reduction in TNF-α and an increase in IL-10 production by PX-treated cells (Figure 4C). To assess the capacity of DCs to activate T cells in response to mycobacterial antigens, we co-cultured DCs and autologous CD4+ T cells from PPD+ donors in the presence or absence of HIF-1α inhibitors and measured the overall IFN-γ and IL-17 production in culture supernatants as well as cell surface expression of T cell markers. We found no significant differences in the activation profile of autologous CD4+ T cells in coculture with iMtb-stimulated DCs treated or not with PX (Figure 4D-E and Figure S4). We conclude that while HIF-1α is important for the maturation of iMtb-stimulated Mo-DCs, it does not influence their capacity to activate CD4+ T cells, at least in vitro.

Revised version: To determine the impact of glycolysis on DC maturation and the capacity to activate T cells, we inhibited HIF-1α activity in iMtb-stimulated DCs employing two HIF-1α inhibitors that display different mechanisms of action. The first is PX-478 (PX) that lowers HIF-1α levels by inhibiting HIF-1α deubiquitination, decreases HIF-1α mRNA expression, and reduces HIF-1α translation37; the second one is echinomycin (Ech), which inhibits the binding of HIF-1α to the hypoxia response element thereby blocking HIF-1α DNA binding capability38, 39. Treatment with either HIF-1α inhibitor PX or Ech diminished lactate release and glucose consumption in iMtb-stimulated DCs without affecting cell viability at the indicated concentration (Figure S3A-F). HIF-1α inhibition by PX significantly abolished ATP production associated with glycolysis without affecting absolute levels of OXPHOS-derived ATP production in iMtb-stimulated DCs (Figure 4A and Figure S3G). In line with these results, the glycolytic capacity was reduced in iMtb-stimulated DCs treated with Ech (Figure S3H). Although HIF-1α inhibitors did not affect the uptake of iMtb by DCs (Figure S3I), we observed a reduction in the expression of activation markers CD83 and CD86, but not in the inhibitory molecule PD-L1, upon treatment with PX (Figure 4B) or Ech (Figure S4). We then measured cytokine production in iMtb-stimulated DCs after HIF-1α inhibition and noted a reduction in TNF-α and an increase in IL-10 production by PX-treated cells (Figure 4C). To assess the capacity of DCs to activate T cells in response to mycobacterial antigens, we co-cultured DCs and autologous CD4+ T cells from PPD+ donors in the presence or absence of HIF-1α inhibitors and measured the overall IFN-γ and IL-17 production in culture supernatants as well as cell surface expression of T cell markers. We found no significant differences in the activation profile of autologous CD4+ T cells in coculture with iMtb-stimulated DCs treated or not with PX (Figure 4D-E and Figure S5). We conclude that while HIF-1α is important for the maturation of iMtb-stimulated Mo-DCs, it does not influence their capacity to activate CD4+ T cells, at least in vitro.

HIF-1α is required for DC maturation upon iMtb stimulation but not for CD4+ T lymphocyte polarization.

Earlier version: (A-C) Mo-DCs were stimulated with irradiated Mtb (iMtb) in the presence or absence of the HIF-1α inhibitor PX-478 (PX). (A) Metabolic flux analysis showing quantification of mitochondrial ATP production and glycolytic ATP production, as in Figure 2C. (B) Mean fluorescence intensity (MFI) of CD83, CD86 and PD-L1 as measured by flow cytometry. (C) TNF-α and IL-10 production by Mo-DCs measured by ELISA. (D-E) Monocytes from PPD+ healthy donors were differentiated towards DCs, challenged or not with iMtb in the presence or absence of PX for 24 h, washed, and co-cultured with autologous CD4+ T cells for 5 days. (D) Extracellular secretion of IFN-γ and IL-17 as measured by ELISA. (E) Absolute abundance of Th1, Th17, Th2 and Th1/Th17 CD4+ T cells after coculture with DCs. When indicated lymphocytes without DCs were cultured (Ly). Statistical significance based on 2-way ANOVA followed by Tukey’s multiple comparison test (∗p < 0.05; ∗∗p < 0.01).

Revised version: (A-C) Mo-DCs were stimulated with irradiated Mtb (iMtb) in the presence or absence of the HIF-1α inhibitor PX-478 (PX). (A) Metabolic flux analysis showing quantification of mitochondrial ATP production and glycolytic ATP production, as in Figure 2C. (B) Mean fluorescence intensity (MFI) of CD83, CD86 and PD-L1 as measured by flow cytometry. (C) TNF-α and IL-10 production by Mo-DCs measured by ELISA. (D-E) Monocytes from PPD+ healthy donors were differentiated towards DCs, challenged or not with iMtb in the presence or absence of PX for 24 h, washed, and co-cultured with autologous CD4+ T cells for 5 days. (D) Extracellular secretion of IFN-γ and IL-17 as measured by ELISA. (E) Absolute abundance of Th1, Th17, Th2 and Th1/Th17 CD4+ T cells after coculture with DCs. When indicated lymphocytes without DCs were cultured (Ly). Statistical significance based on 2-way ANOVA followed by Tukey’s multiple comparison test (∗p < 0.05; ∗∗p < 0.01). The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown.

HIF-1α-mediated glycolysis triggers the motility of DCs upon iMtb stimulation

Earlier version: Since DC migration to lymph nodes is essential to initiate an adaptive immune response and glycolytic activity has been reported to control DC migration upon stimulation22,23, we evaluated the migratory properties of iMtb-stimulated DCs in the presence of inhibitors of HIF-1α and glycolysis. First, we confirmed that PX and oxamate (OX), a well-established glycolysis inhibitor, diminished the glycolytic activity of iMtb-stimulated human Mo-DCs, as demonstrated by reduced lactate release (Figure S2A and 5A). Next, using a transwell migration assay, we found that PX and OX treatment significantly diminished the chemotactic activity of iMtb-stimulated human Mo-DCs in response to CCL21 (Figure 5B), a CCR7 ligand responsible for the migration of DCs into lymphoid organs. We also assessed the 3 dimentional (3D) migration capacity of iMtb-stimulated DCs through a collagen matrix in which DCs use an amoeboid migration mode41 and found that 3D migration was significantly impaired upon HIF-1α or glycolysis inhibition (Figure 5C). The role of glycolysis in the migration of iMtb-stimulated Mo-DCs was further confirmed using an additional glycolysis inhibitor, GSK2837808A, which reduced both the release of lactate by iMtb-stimulated Mo-DCs and their migration in response to CCL21 (Figure S5A-B). To further investigate the effects of glycolysis on cell migration, we turned to an in vivo model. Murine bone marrow-derived DCs (BMDCs) isolated and stimulated with iMtb in the presence or absence of PX or OX were labeled with CFSE and transferred into naïve mice (Figure 5D). Similarly to human Mo-DCs, iMtb stimulation increased glycolysis in BMDCs, which was inhibited by PX and OX treatment in vitro (Figure S5C). Three hours after the transfer of BMDCs into recipient mice, nearby lymph nodes were collected for DC quantification (Figure 5D). A higher number of adoptively transferred DCs (CFSE-labeled CD11c+ cells) were detected in lymph nodes from mice that received iMtb-stimulated BMDCs compared to mice that received untreated BMDCs or iMtb-BMDCs treated with either PX or OX (Figure 5E and Figure S5D). Of note, we verified that CCR7 expression on iMtb-stimulated BMDCs was not affected by OX or PX treatment, so the effect could not be ascribed to downregulation of the chemokine receptor (Figure S5E). Therefore, we conclude that HIF-1α-mediated glycolysis is required for the successful migration of iMtb-stimulated DCs into lymph nodes.

Revised version: Since DC migration to lymph nodes is essential to initiate an adaptive immune response and glycolytic activity has been reported to control DC migration upon stimulation22, 23, we evaluated the migratory properties of iMtb-stimulated DCs in the presence of inhibitors of HIF-1α and glycolysis. First, we confirmed that PX and oxamate (OX), a well-established glycolysis inhibitor, diminished the glycolytic activity of iMtb-stimulated human Mo-DCs, as demonstrated by reduced lactate release (Figure S3A and 5A). Next, using a transwell migration assay, we found that PX and OX treatment significantly diminished the chemotactic activity of iMtb-stimulated human Mo-DCs in response to CCL21 (Figure 5B), a CCR7 ligand responsible for the migration of DCs into lymphoid organs. We also assessed the 3 dimensional (3D) migration capacity of iMtb-stimulated DCs through a collagen matrix in which DCs use an amoeboid migration mode40 and found that 3D migration was significantly impaired upon HIF-1α or glycolysis inhibition (Figure 5C). The role of glycolysis in the migration of iMtb-stimulated Mo-DCs was further confirmed using an additional glycolysis inhibitor, GSK2837808A, which reduced both the release of lactate by iMtb-stimulated Mo-DCs and their migration in response to CCL21 (Figure S6A-B). To further investigate the effects of glycolysis on cell migration, we turned to an in vivo model. Murine bone marrow-derived DCs (BMDCs) isolated and stimulated with iMtb in the presence or absence of PX or OX were labeled with CFSE and transferred into naïve mice (Figure 5D). Similar to human Mo-DCs, iMtb stimulation increased glycolysis in BMDCs, which was inhibited by PX and OX treatment in vitro (Figure S6C). Three hours after the transfer of BMDCs into recipient mice, nearby lymph nodes were collected for DC quantification (Figure 5D). A higher number of adoptively transferred DCs (CFSE-labeled CD11c+ cells) were detected in lymph nodes from mice that received iMtb-stimulated BMDCs compared to mice that received untreated BMDCs or iMtb-BMDCs treated with either PX or OX (Figure 5E and Figure S6D). Of note, we verified that CCR7 expression on iMtb-stimulated BMDCs was not affected by OX or PX treatment, so the effect could not be ascribed to downregulation of the chemokine receptor (Figure S6E). Therefore, we conclude that HIF-1α-mediated glycolysis is required for the successful migration of iMtb-stimulated DCs into lymph nodes.

HIF1α-mediated-glycolysis is required to trigger migratory activity in iMtb-stimulated DCs.

Earlier version: Mo-DCs were treated or not with HIF-1a inhibitor PX-478 (PX) or glycolysis inhibitor oxamate (OX) and stimulated with iMtb for 24h. (A) Lactate release as measured in supernatants in DCs stimulated or not with iMtb in the presence of OX. (B) Percentage of migrated cells towards CCL21 relative to the number of initial cells per condition. (C) Three-dimensional amoeboid migration of DCs through a collagen matrix after 24h. Cells within the matrix were fixed and stained with DAPI. Images of the membrane of each insert were taken and the percentage of cells per field were counted. (D) Representative schematic of the experimental setup for in-vivo migration assays. (E) Percentages of migrating BMDCs (CFSE-labeled among CD11c+) recovered from inguinal lymph nodes. Statistical significance assessed by (A-B) ANOVA followed by Dunnett’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (C) Nested ANOVA followed by Dunnett’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (E) ANOVA followed by Holm-Sidak’s multiple comparisons test (∗p < 0.05).

Revised version: Mo-DCs were treated (or not) with HIF-1a inhibitor PX-478 (PX) or glycolysis inhibitor oxamate (OX) and stimulated with iMtb for 24 h. (A) Lactate release as measured in supernatants in DCs stimulated or not with iMtb in the presence of OX. (B) Percentage of migrated cells towards CCL21 relative to the number of initial cells per condition. (C) Three-dimensional amoeboid migration of DCs through a collagen matrix after 24 h. Cells within the matrix were fixed and stained with DAPI. Images of the membrane of each insert were taken and the percentage of cells per field were counted. The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown. (D) Representative schematic of the experimental setup for in vivo migration assays. (E) Percentages of migrating BMDCs (CFSE-labeled among CD11c+) recovered from inguinal lymph nodes. Statistical significance assessed by (A-B) ANOVA followed by Dunnett’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (C) Nested ANOVA followed by Dunnett’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (E) ANOVA followed by Holm-Sidak’s multiple comparisons test (∗p < 0.05).

Stabilization of HIF-1α promotes migration of tolerogenic DCs and DCs derived from TB patient monocytes

Earlier version: Since DC differentiation is skewed, at least partially, towards a tolerogenic phenotype during TB42–44, we investigated whether tolerogenic DCs can be reprogrammed into immunogenic DCs by modulating their glycolytic pathway after iMtb stimulation. To this end, we generated tolerogenic Mo-DCs by adding dexamethasone (Dx) before stimulation with iMtb in the presence or absence of dimethyloxalylglycine (DMOG), which stabilizes the expression of HIF-1α. HIF-1α expression is tightly regulated by prolyl hydroxylase domain containing proteins which facilitate the recruitment of the von Hippel-Lindau (VHL) protein, leading to ubiquitination and degradation of HIF-1α by the proteasomes45. DMOG inhibits the prolyl hydroxylase domain-containing proteins. Acquisition of the tolerogenic phenotype was confirmed by the lack of upregulation of costimulatory markers CD83 and CD86, as well as by increased PD-L1 expression in iMtb-DCs treated with Dx compared to control iMtb-DCs (Figure S6A). Moreover, Dx-treated DCs did not exhibit an increase in lactate release, consumption of glucose or induction of HIF-1α expression in response to iMtb, showing a high consumption of levels of glucose under basal conditions (Figure 6A-B). Of note, HIF-1α stabilization using DMOG restored the HIF-1α expression and lactate production in response to iMtb in Dx-treated DCs and increased the consumption of glucose (Figure 6A-B). Activation of HIF-1α also improved 3D amoeboid migration, as well as 2D migration capacity of DCs towards CCL21 of iMtb-stimulated Dx-treated DCs (Figure 6C-D and Figure S6B). Confirming the relevance of these findings to human TB patients, we found that iMtb-stimulated Mo-DCs from TB patients were deficient in their capacity to migrate towards CCL21 (Figure 6E) and in glycolytic activity, compared to Mo-DCs from healthy subjects (Figure 6F-G). Strikingly, stabilizing HIF-1α expression using DMOG in Mo-DCs from TB patients restored their chemotactic activity in response to iMtb (Figure 6H). These data indicate that the impaired migratory capacity of iMtb-stimulated tolerogenic DCs or TB patient-derived DCs can be restored via HIF-1α stabilization; thus, glycolysis is critical for DC function during TB in both murine and human contexts.

Revised version: Since DC differentiation is skewed, at least partially, towards a tolerogenic phenotype during TB41–43 we investigated whether tolerogenic DCs can be reprogrammed into immunogenic DCs by modulating their glycolytic pathway after iMtb stimulation. To this end, we generated tolerogenic Mo-DCs by adding dexamethasone (Dx) before stimulation with iMtb in the presence or absence of dimethyloxalylglycine (DMOG), which stabilizes the expression of HIF-1α. HIF-1α expression is tightly regulated by prolyl hydroxylase domain containing proteins which facilitate the recruitment of the von Hippel-Lindau (VHL) protein, leading to ubiquitination and degradation of HIF-1α by the proteasomes44. DMOG inhibits the prolyl hydroxylase domain-containing proteins. Acquisition of the tolerogenic phenotype was confirmed by the lack of upregulation of costimulatory markers CD83 and CD86, as well as by increased PD-L1 expression in iMtb-DCs treated with Dx compared to control iMtb-DCs (Figure S7A). Moreover, Dx-treated DCs did not exhibit an increase in lactate release, consumption of glucose or induction of HIF-1α expression in response to iMtb, showing a high consumption of levels of glucose under basal conditions (Figure 6A-B). Of note, HIF-1α stabilization using DMOG restored the HIF-1α expression and lactate production in response to iMtb in Dx-treated DCs and increased the consumption of glucose (Figure 6A-B). Activation of HIF-1α also improved 3D amoeboid migration, as well as 2D migration capacity of DCs towards CCL21 of iMtb-stimulated Dx-treated DCs (Figure 6C-D and Figure S7B). Confirming the relevance of these findings to human TB patients, we found that iMtb-stimulated Mo-DCs from TB patients were deficient in their capacity to migrate towards CCL21 (Figure 6E) and in glycolytic activity, compared to Mo-DCs from healthy subjects (Figure 6F-G). Strikingly, stabilizing HIF-1α expression using DMOG in Mo-DCs from TB patients restored their chemotactic activity in response to iMtb (Figure 6H). These data indicate that the impaired migratory capacity of iMtb-stimulated tolerogenic DCs or TB patient-derived DCs can be restored via HIF-1α stabilization; thus, glycolysis is critical for DC function during TB in both murine and human contexts.

Stabilization of HIF-1α promotes migration of tolerogenic DCs and Mo-DCs from TB patients.

Earlier version: Tolerogenic Mo-DCs were generated by dexamethasone (Dx) treatment and were stimulated or not with iMtb in the presence or absence of HIF-1α activator DMOG. (A) Lactate release and glucose uptake as measured in supernatant. (B) Mean fluorescence intensity (MFI) of HIF-1α. Representative histograms and quantification are shown. (C) Three-dimensional amoeboid migration of DCs through a collagen matrix. After 24h of migration, images of stacks within the matrix were taken every 30 µm. Percentage of migrating cells were defined as cells in the stacks within the matrix relative to total number of cells. (D) Chemotactic activity towards CCL21 in vitro. (E-H) Mo-DCs were generated from healthy subjects (HS) or TB patients, and DCs were stimulated or not with iMtb. (E) Chemotaxis index towards CCL21 (relative to unstimulated DCs). (F) Lactate production ratio relative to unstimulated DCs. (G) Glycolytic capacity assessed by SCENITH. (H) Chemotactic activity towards CCL21 of Mo-DCs from TB patients stimulated with iMtb and treated or not with DMOG. Statistical significance assessed by (A-D) 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (E-G) Unpaired T test (∗p < 0.05); (H) Paired T test (∗p < 0.05).

Revised version: Tolerogenic Mo-DCs were generated by dexamethasone (Dx) treatment and were stimulated (or not) with iMtb in the presence or absence of HIF-1α activator DMOG. (A) Lactate release and glucose uptake as measured in supernatant. (B) Mean fluorescence intensity (MFI) of HIF-1α. Representative histograms and quantification are shown. (C) Three-dimensional amoeboid migration of DCs through a collagen matrix. After 24 h of migration, images of stacks within the matrix were taken every 30 µm. Percentage of migrating cells were defined as cells in the stacks within the matrix relative to total number of cells. (D) Chemotactic activity towards CCL21 in vitro. (E-H) Mo-DCs were generated from healthy subjects (HS) or TB patients, and DCs were stimulated (or not) with iMtb. (E) Chemotaxis index towards CCL21 (relative to unstimulated DCs). (F) Lactate production ratio relative to unstimulated DCs. (G) Glycolytic capacity assessed by SCENITH. (H) Chemotactic activity towards CCL21 of Mo-DCs from TB patients stimulated with iMtb and treated or not with DMOG. Statistical significance assessed by (A-D) 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01); (E-G) Unpaired T test (∗p < 0.05); (H) Paired T test (∗p < 0.05). The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown.

Earlier version: Glycolytic activation in CD16+ monocytes from TB patients leads to DCs with poor migration capacity

Revised version: CD16+ monocytes from TB patients show increased glycolytic capacity

Added in the revised version: Since we observed differences in the metabolic activity of DCs derived from monocytes of TB patients when compared to healthy donors, we next focused on evaluating the release of lactate by DC precursors from both subject groups during the first hours of DC differentiation with IL-4/GM-CSF. We found a high release of lactate by monocytes from TB patients compared to healthy donors after 1 h of differentiation (Figure 7A). Lactate accumulation increased in both subject groups after 24 h with IL-4/GM-CSF (Figure 7A). Based on these differential glycolytic activities displayed by DC precursors from both subject groups at very early stages of the differentiation process, we decided to evaluate the ex vivo metabolic profile of monocytes using SCENITH. To this end, we assessed the baseline glycolytic capacity of the three main populations of monocytes: classical (CD14+CD16−), intermediate (CD14+CD16+), and non-classical (CD14dimCD16+) monocytes. We found that both populations of CD16+ monocytes from TB patients had a higher glycolytic capacity than monocytes from healthy donors (Figure 7B). Moreover, the glycolytic capacity of CD16+ monocytes (CD14+CD16+ and CD14dimCD16+) correlates with time since the onset of TB-related symptoms (Figure 7C), with no association to the extent or severity of lung disease (unilateral/bilateral lesions and with/without cavities, Figure S8). To further expand the metabolic characterisation of monocyte subsets from TB patients, we used previously published transcriptomic data (GEO accession number: GSE185372) of CD14+CD16-, CD14+CD16+ and CD14dimCD16+ monocytes isolated from individuals with active TB, latent TB (IGRA+), as well as from TB negative healthy controls (IGRA-)45. Within this framework, we performed high-throughput GeneSet Enrichment Analysis (GSEA) using the BubbleMap module of BubbleGUM, which includes a multiple testing correction step to allow comparisons between the three monocyte subsets46. As expected, this approach reveals enrichments in genes associated with interferon responses (alpha and gamma) in patients with active TB compared to healthy donors (either IGRA- or latent TB) for all three monocyte subsets (Figure 7D). Consistent with our findings, glycolysis increases in active TB in both CD14+CD16+ and CD14dimCD16+ monocytes (albeit not significant), while it appears to decrease in classical CD14+CD16- monocytes (Figure 7D). Unlike CD14+CD16- cells, the inflammatory response is notably enriched in CD14+CD16+ and CD14dimCD16+ monocytes from patients with active TB compared those with latent TB or healthy subjects (Figure 7D), suggesting that their glycolytic profile correlates with a higher inflammatory state. Finally, no significant enrichment of oxidative phosphorylation-associated genes was found in any of the performed comparisons (Figure 7D). Taken together, these results demonstrate that TB disease is associated with an increased activation and glycolytic profile of circulating CD16+ monocytes.

Added in the revised version: CD16+ monocytes from TB patients show increased glycolytic capacity.

Removed in the revised version: Since we observed differences in the metabolic activity of DCs derived from monocytes of TB patients when compared to healthy donors, we next focused on evaluating the release of lactate by DC precursors from both subject groups during the first hours of DC differentiation with IL-4/GM-CSF. We found a high release of lactate by monocytes from TB patients compared to healthy donors after 1 h of differentiation (Figure 7A). Lactate accumulation increased in both subject groups after 24h with IL-4/GM-CSF (Figure 7A). Based on these differential glycolytic activities displayed by DC precursors from both subject groups at very early stages of the differentiation process, we decided to evaluate the ex vivo metabolic profile of monocytes using SCENITH. To this end, we assessed the baseline glycolytic capacity of the three main populations of monocytes: classical (CD14+, CD16−), intermediate (CD14+, CD16+), and non-classical (CD14dim, CD16+) monocytes. We found that both populations of CD16+ monocytes from TB patients had a higher glycolytic capacity than monocytes from healthy donors (Figure 7B). As we have previously demonstrated that CD16+ monocytes from TB patients generate aberrant DCs42, we hypothesized that the different metabolic profile of this monocyte subset could yield DCs with exhausted glycolytic activity and thus lower migration capacity upon Mtb exposure. To test this hypothesis, we treated with DMOG to increase the activity of HIF-1α during the first 24h of monocyte differentiation from healthy donors, leading to an exacerbated increase in lactate release at early stages of the differentiation (Figure 7C). Such early addition of DMOG to healthy monocytes resulted in the generation of DCs (6 days with IL-4/GM-CSF) characterized by equivalent levels of CD1a as control DCs, with a significant decrease in the expression of DC-SIGN (Figure 7D). In terms of activation marker expression, DCs differentiated from DMOG-pretreated cells responded to iMtb by upregulating CD86 at higher levels compared to control cells, with an accompanying trend towards reduced upregulation of CD83 (Figure 7E). We also observed that DCs from DMOG pretreated-cells exhibited a lower migratory capacity in response to iMtb (Figure 7F), reminiscent of the 2D migration capacities of Mo-DCs from TB patients. Altogether, our data suggest that the activated glycolytic status of pathological monocytes from TB patients leads to the generation of DCs with low motility in response to Mtb.

Removed in the revised version: Exacerbation of glycolysis in CD16+ monocytes from TB patients leads to DCs with poor migration capacity.

Earlier version: (A) Monocytes from TB patients or healthy subjects (HS) were isolated and cultured with IL-4 and GM-CSF for 24h. Accumulation of lactate in culture supernatants were measured at 1 and 24h of differentiation. (B) Glycolytic capacity measured by SCENITH of monocyte subsets as defined by their CD14 and CD16 expression from HS and TB patients. (C-F) Monocytes from HS were treated with DMOG during the first 24h of differentiation with IL-4/GM-CSF (earlyDMOG) and removed afterwards. On day 6 of differentiation, cells were stimulated or not with iMtb. (C) Monocyte lactate release after 24h of DMOG addition. (D) Percentage of cells expressing CD1a and DC-SIGN after full 6 days of differentiation (E) MFI of CD83 and CD86 on day 6 differentiated cells stimulated or not with iMtb for further 24h. (F) Chemotactic activity towards CCL21 of DCs. Statistical significance was assessed by (A) Paired T test for 0 vs. 24h (∗p < 0.05) and 2-way ANOVA for HS vs. TB at each time (∗∗p < 0.01); (B) unpaired T test (∗p < 0.05; ∗∗p < 0.01); (C-D) paired T test (∗∗p < 0.01); (E-F) 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05; ∗∗p < 0.01).

Revised version: (A) Monocytes from TB patients or healthy subjects (HS) were isolated and cultured with IL-4 and GM-CSF for 24 h. Accumulation of lactate in culture supernatants were measured at 1 and 24 h of differentiation. (B) Glycolytic capacity measured by SCENITH of monocyte subsets as defined by their CD14 and CD16 expression from HS and TB patients. (C) Correlation analysis between the baseline glycolytic capacity and the evolution time of TB symptoms for each monocyte subset (CD14+CD16-, CD14+CD16+ and CD14dimCD16+, n = 14). Linear regression lines are shown. Spearman’s rank test. The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown. (D) BubbleMap analysis, a high-throughput extension of GSEA, on the pairwise comparisons of monocytes from healthy patients (HS) or donors with latent TB (LTB) vs patients with active TB (TB), for each monocyte subset (CD14+CD16-, CD14+CD16+ and CD14dimCD16+). The gene sets shown come from the Hallmark (H.) collection of the Molecular Signature Database (MSigDB). Statistical significance was assessed by (A) Paired T test for 0 vs. 24h (∗p < 0.05) and 2-way ANOVA for HS vs. TB at each time (∗∗p < 0.01); (B) unpaired T test (∗p < 0.05; ∗∗p < 0.01).

Added in the revised version: HIF1-α activation in CD16+ monocytes from TB patients leads to differentiated DCs with a poor migration capacity

Added in the revised version: Since circulating CD16+ monocytes from TB patients are highly glycolytic, we evaluated the expression of HIF1-α among the populations. We found that CD16+ monocytes from TB patients exhibited a higher expression of HIF1-α than from healthy donors (Figure 8A). As we have previously demonstrated that CD16+ monocytes from TB patients generate aberrant DCs41, we hypothesized that the different metabolic profile of this monocyte subset could yield DCs with some sort of exhausted glycolytic capacity and thus lower migration activity upon Mtb exposure. To test this hypothesis, we treated with DMOG to increase the activity of HIF-1α during the first 24 h of monocyte differentiation from healthy donors, leading to an exacerbated increase in lactate release at early stages of the differentiation (Figure 8B). Such early addition of DMOG to healthy monocytes resulted in the generation of DCs (6 days with IL-4/GM-CSF) characterized by equivalent levels of CD1a as control DCs, with a significant decrease in the expression of DC-SIGN (Figure S9A). In terms of activation marker expression, DCs differentiated from DMOG-pretreated cells responded to iMtb by upregulating CD86 at higher levels compared to control cells, with an accompanying trend towards reduced upregulation of CD83 (Figure S9B). We also observed that DCs from DMOG pretreated-cells exhibited a lower migratory capacity in response to iMtb (Figure 8C), reminiscent of the 2D migration capacities of Mo-DCs from TB patients. Altogether, our data suggest that the activated glycolytic status of monocytes from TB patients leads to the generation of DCs with low motility in response to Mtb.

Added in the revised version: HIF1-α activation in CD16+ monocytes from TB patients leads to DCs with poor migration capacity.

Added in the revised version: (A) Ex-vivo determination of HIF-1α expression by monocytes from healthy subjects (HS) or TB patients (TB) for each monocyte subset (CD14+CD16-, CD14+CD16+ and CD14dimCD16+). (B-C) Monocytes from HS were treated with DMOG during the first 24 h of differentiation with IL-4/GM-CSF (earlyDMOG) and removed afterwards. On day 6 of differentiation, cells were stimulated (or not) with iMtb. (B) Monocyte lactate release after 24 h of DMOG addition. (C) Chemotactic activity towards CCL21 of DCs. Statistical significance was assessed by (B) paired T test (∗∗p < 0.01); (C) 2-way ANOVA followed by Tukey’s multiple comparisons test (∗p < 0.05). The data are represented as scatter plots with each circle representing a single individual, means ± SEM are shown.

Discussion

Earlier version: In this study, we provide evidence for the role of HIF-1α-mediated glycolysis in promoting the migratory capacity of DCs upon encounter with Mtb. Our approach to quantify the ex vivo metabolism of monocytes shows that CD16+ monocytes from TB patients display a pathological glycolytic activity that may result in the generation of DCs with poor migratory capacities in response to Mtb. Our results suggest that under extensive chronic inflammatory conditions, such as those found in TB patients, circulating monocytes may be metabolically preconditioned to differentiate into DCs with low immunogenic potential.

Revised version: In this study, we provide evidence for the role of HIF-1α-mediated glycolysis in promoting the migratory capacity of DCs upon encounter with iMtb. Our approach to quantify the ex vivo metabolism of monocytes shows that CD16+ monocytes from TB patients display an exacerbated glycolytic activity that may result in the generation of DCs with poor migratory capacities in response to iMtb. Our results suggest that under extensive chronic inflammatory conditions, such as those found in TB patients, circulating monocytes may be metabolically preconditioned to differentiate into DCs with low migratory potential.

Earlier version: Upon Mtb infection of naive mice, initial accumulation of activated CD4+ T cells in the lung is delayed, occurring between 2-3 weeks postinfection3,46. The absence of sterilizing immunity induced by TB vaccines (such as BCG) has been proposed to result from delayed activation of DCs and the resulting delay in antigen presentation and activation of vaccine-induced CD4+ T-cell responses47. In this context, it was demonstrated that Mtb-infected Mo-DCs recruited to the site of infection exhibit low CCR7 expression and impaired migration to lymph nodes compared to uninfected Mo-DCs48. Additionally, Mo-DCs have been found to play a key role in transporting Mtb antigens from the lung to the draining lymph node, where conventional DCs present antigens to naive T cells49. The migratory capacity of responding DCs is thus of paramount importance to the host response to Mtb infection.

Revised version: Upon Mtb infection of naive mice, initial accumulation of activated CD4+ T cells in the lung is delayed, occurring between 2-3 weeks post-infection3, 47. The absence of sterilizing immunity induced by TB vaccines, such as BCG, has been proposed to result from delayed activation of DCs and the resulting delay in antigen presentation and activation of vaccine-induced CD4+ T-cell responses48. In this context, it was demonstrated that Mtb-infected Mo-DCs recruited to the site of infection exhibit low CCR7 expression and impaired migration to lymph nodes compared to uninfected Mo-DCs49. Additionally, Mo-DCs have been found to play a key role in transporting Mtb antigens from the lung to the draining lymph node, where conventional DCs present antigens to naive T cells50. The migratory capacity of responding DCs is thus of paramount importance to the host response to Mtb infection.

Earlier version: In this study, we found that Mtb exposure triggers glycolysis in Mo-DCs, which promotes their migration capacity in a HIF-1α-dependent manner. Recently, it was shown that glycolysis was required for CCR7-triggered murine DC migration in response to LPS22–24. Glycolysis was also reported to be required for the migration of other immune cells such as macrophages50 and regulatory T cells51. Consistently, we show that inhibition of HIF-1α-dependent glycolysis impairs human Mo-DC migration upon Mtb stimulation. The linkage between cellular metabolism and migratory behavior are supported by studies that have elucidated how glycolysis can be mechanically regulated by changes in the architecture of the cytoskeleton, ultimately impacting the activity of glycolytic enzymes52,53. In addition, interesting links between cellular mechanics and metabolism have been previously described for DCs, highlighting the potential to alter DC mechanics to control DC trafficking and consequently T cell priming16. However, studies focused on the molecular mechanisms by which metabolic pathways impact the machinery responsible for cell movement in the context of TB infection will be required to more fully understand potential handles for therapeutic manipulation.

Revised version: In this study, we found that Mtb exposure triggers glycolysis in Mo-DCs from healthy donors, which promotes their migration capacity in a HIF-1α-dependent manner. Recently, it was shown that glycolysis was required for CCR7-triggered murine DC migration in response to LPS22–24. Glycolysis was also reported to be required for the migration of other immune cells such as macrophages51 and regulatory T cells52. Consistently, we show that inhibition of HIF-1α-dependent glycolysis impairs human Mo-DC migration upon Mtb stimulation. The link between cellular metabolism and migratory behavior are supported by studies that have elucidated how glycolysis can be mechanically regulated by changes in the architecture of the cytoskeleton, ultimately impacting the activity of glycolytic enzymes53, 54. In addition, interesting links between cellular mechanics and metabolism have been previously described for DCs, highlighting the potential to alter DC mechanics to control DC trafficking and consequently T cell priming16. However, studies focused on the molecular mechanisms by which metabolic pathways impact the machinery responsible for cell movement in the context of TB infection will be required to better understand and design therapeutic manipulation.

Earlier version: We demonstrated that tolerogenic DCs induced by dexamethasone as well as DCs derived from TB patient monocytes exhibit lower lactate release and impaired trafficking toward CCL21 upon Mtb stimulation; both of these phenotypes could be rescued by stabilization of HIF-1α expression. To our knowledge, this is the first study addressing the metabolic status and migratory activity of Mo-DCs from TB patients. With respect to the metabolism of tolerogenic DCs broadly, our results are consistent with reported data showing that DC tolerance can be induced by drugs promoting OXPHOS, such as vitamin D and dexamethasone54–56. It was interesting to note that, although migration of tolerogenic DCs did not increase upon Mtb stimulation, it was increased under basal conditions, which is in agreement with previous data showing a high steady-state migration capacity of putatively tolerogenic DCs57.

Revised version: We demonstrated that tolerogenic DCs induced by dexamethasone as well as DCs derived from TB patient monocytes exhibit lower lactate release and impaired trafficking toward CCL21 upon Mtb stimulation; both phenotypes could be rescued by stabilization of HIF-1α expression. To our knowledge, this is the first study to address how the metabolic status of monocytes from TB patients influences the migratory activity of further differentiated DCs. According to our findings, the activation status of the glycolysis/HIF-1α axis in monocytes would be a predictor of refractoriness to differentiation into migratory DCs in TB. With respect to the metabolism of tolerogenic DCs broadly, our results are consistent with reported data showing that DC tolerance can be induced by drugs promoting OXPHOS, such as vitamin D and dexamethasone55–57. It was interesting to note that, although migration of tolerogenic DCs did not increase upon Mtb stimulation, it was increased under basal conditions, which agrees with previous data showing a high steady-state migration capacity of putatively tolerogenic DCs58.

Earlier version: It has been widely demonstrated that immune cells have the ability to switch to glycolysis following engagement of TLRs58. Our work showed that TLR2 ligation by either viable or irradiated Mtb was necessary to trigger glycolysis in DCs, at least at early times post-stimulation. It remains to be elucidated whether persistent interaction between DCs and Mtb might lead to an attenuation in glycolysis over time, as has been reported for macrophages59. In this regard, our data demonstrates that chronic Mtb infection leads to monocytes bearing an exacerbated glycolytic status which results in DCs with exhausted glycolytic capacity. Although DCs stimulated with iMtb in the presence of HIF-1α inhibitor exhibited differences in activation markers and cytokine profile, we found that they were still able to activate CD4+ T cells from PPD+ donors in response to iMtb. These findings complement previous evidence showing that LPS-induced mature DCs inhibit T-cell responses through HIF-1α activation in the presence of glucose, leading to greater T cell activation capacity in low glucose contexts such as at the synapsis between DCs and T cells25. In this work, we did not detect an impact on T cell activation upon HIF-1α inhibition in DCs, but we observed a clear reduction in their migration capacity that may limit or delay DC encounters with T cells in vivo, leading to poor T cell activation in the lymph nodes. In this regard, mouse studies have shown that DC migration directly correlates with T cell proliferation60. However, we cannot rule out the possibility that other CD4+ T cell subsets (such as regulatory T cells), CD1-restricted T cells, and/or CD8+ T cell subsets could be differentially activated by iMtb-stimulated DCs lacking HIF-1α activity.

Revised version: It has been widely demonstrated that immune cells can switch to glycolysis following engagement of TLRs59. Our work showed that TLR2 ligation by either viable or irradiated Mtb was necessary to trigger glycolysis in DCs, at least at early times post-stimulation. It remains to be elucidated whether persistent interaction between DCs and Mtb might lead to an attenuation in glycolysis over time, as has been reported for macrophages60. In this regard, our data demonstrates that chronic Mtb infection leads to monocytes bearing an exacerbated glycolytic status likely tied to prolonged and or excessive stimulation of membrane bound TLRs in circulation, which results in DCs with some sort of exhausted glycolytic capacity. Although DCs stimulated with iMtb in the presence of a HIF-1α inhibitor exhibited differences in activation markers and cytokine profile, we found that they were still able to activate CD4+ T cells from PPD+ donors in response to iMtb. These findings complement previous evidence showing that LPS-induced mature DCs inhibit T-cell responses through HIF-1α activation in the presence of glucose, leading to greater T cell activation capacity in low glucose contexts such as at the interface between DCs and T cells25. In this work, we did not detect an impact on T cell activation upon HIF-1α inhibition in DCs, but we observed a clear reduction in their migration capacity that may limit or delay DC encounters with T cells in vivo, leading to poor T cell activation in the lymph nodes. In this regard, mouse studies have shown that DC migration directly correlates with T cell proliferation61. However, we cannot rule out the possibility that other CD4+ T cell subsets (such as regulatory T cells), CD1-restricted T cells, and/or CD8+ T cell subsets could be differentially activated by iMtb-stimulated DCs lacking HIF-1α activity.

Earlier version: Three different populations of human monocytes have been identified: classical (CD14+, CD16−), intermediate (CD14+, CD16+), and non-classical (CD14dim, CD16+) monocytes61. These monocyte subsets are phenotypically and functionally distinct. Classical monocytes readily extravasate into tissues in response to inflammation, where they can differentiate into macrophage-like or DC-like cells62; intermediate monocytes are well-suited for antigen presentation, cytokine secretion, and differentiation; and non-classical monocytes are involved in complement and Fc gamma-mediated phagocytosis and their main function is cell adhesion63,64. Unlike non-classical monocytes, the two CD14+ monocyte populations are known to extravasate into tissues and thus are likely to act as precursors capable at giving rise to Mo-DCs in inflamed tissues. However, the DC differentiation capacity of the intermediate population is still not well defined. We previously demonstrated that monocytes from TB patients generate aberrant DCs, and that CD16+ monocytes generate aberrant DCs upon treatment with GM-CSF and IL-442. Here, we demonstrated that glycolysis seems to play a dual role during DC differentiation from monocytes, on the one side, being required for fully differentiated-DC migration to lymph nodes in response to Mtb and, on the other side, leading to DCs with poor iMtb-responsive migratory capacity if activated during the onset of DC differentiation. Similarly, we found that CD16+ cells from TB patients display a pathologically activated glycolytic status. Additionally, we showed that monocytes from TB patients are not only enriched in CD16+ cells, but also display an altered chemokine receptor expression profile65, demonstrating that both phenotype and function of a given monocyte subset may differ under pathological conditions. While it is difficult to determine whether the heightened glycolytic profile of monocytes may limit their differentiation into DCs in vivo, we provided evidence that an increase in HIF1α-mediated glycolysis in precursors leads to the generation of cells with poor ability to migrate in response to CCL21 in vitro. In line with this observation, a recent paper revealed a significant increase in the glycolytic capacity occurs during the first 24h of monocyte differentiation towards a tolerogenic DC phenotype, as induced by vitamin D326, highlighting the detrimental role of an early activated inflammatory profile in DC precursors. A possible explanation for these effects may be found in lactate accumulation in monocytes during DC differentiation. Lactate signaling in immune cells leads to metabolic alterations in DCs that program them to a regulatory state66, and lactate has also been shown to suppress DC differentiation and maturation18; thus, excessive precursor glycolytic activity may result in DCs biased toward regulatory functions.

Revised version: Three different populations of human monocytes have been identified: classical (CD14+, CD16−), intermediate (CD14+, CD16+), and non-classical (CD14dim, CD16+) monocytes62. These monocyte subsets are phenotypically and functionally distinct. Classical monocytes readily extravasate into tissues in response to inflammation, where they can differentiate into macrophage-like or DC-like cells63; intermediate monocytes are well-suited for antigen presentation, cytokine secretion, and differentiation; and non-classical monocytes are involved in complement and Fc gamma-mediated phagocytosis and their main function is cell adhesion64, 65. Unlike non-classical monocytes, the two CD14+ monocyte populations are known to extravasate into tissues and thus are likely to act as precursors capable of giving rise to Mo-DCs in inflamed tissues. However, the DC differentiation capacity of the intermediate population is still not well defined. We previously demonstrated that monocytes from TB patients generate aberrant DCs, and that CD16+ monocytes generate aberrant DCs upon treatment with GM-CSF and IL-441. Here, we demonstrated that glycolysis seems to play a dual role during DC differentiation from monocytes, on the one side, being required for fully differentiated-DC migration to lymph nodes in response to Mtb and, on the other side, leading to DCs with poor iMtb-responsive migratory capacity if activated during the onset of DC differentiation. In this regard, DCs from healthy subjects respond to iMtb by inducing a glycolytic and migratory profile, while monocytes isolated from TB patients exhibit an unusual early glycolytic state that results in the ulterior generation of DCs with low glycolytic and migratory activities in response to Mtb. Similarly, we found that CD16+ cells from TB patients display an activated glycolytic status, as well as elevated HIF-1α expression levels compared to their healthy counterparts. Additionally, we showed that monocytes from TB patients are not only enriched in CD16+ cells, but also display an altered chemokine receptor expression profile66, demonstrating that both phenotype and function of a given monocyte subset may differ under pathological conditions. While it is difficult to determine whether the heightened glycolytic profile of monocytes may limit their differentiation into DCs in vivo, we provided evidence that an increase in HIF1α-mediated glycolysis in precursors leads to the generation of cells with poor ability to migrate in response to CCL21 in vitro. In line with this observation, a recent study revealed a significant increase in the glycolytic capacity occurs during the first 24 h of monocyte differentiation towards a tolerogenic DC phenotype, as induced by vitamin D326, highlighting the detrimental role of an early activated inflammatory profile in DC precursors. A possible explanation for these effects may be found in lactate accumulation in monocytes during DC differentiation. Lactate signaling in immune cells leads to metabolic alterations in DCs that program them to a regulatory state67, and lactate has also been shown to suppress DC differentiation and maturation18; thus, excessive precursor glycolytic activity may result in DCs biased toward regulatory functions.

Earlier version: Taken together, our data provide novel insights into the immunometabolic pathways involved in the trafficking of DCs to the lymph nodes, which may be targetable to generate sterilizing vaccine-induced immunity against TB. To this end, we propose that the development of tools aiming to improve the migratory capacity of DCs by fostering their HIF-1α-mediated glycolytic activity may increase the effectiveness of preventive strategies for TB, especially given that the number of DCs to reach the lymph node was shown to be a critical parameter for the outcome of DC-based vaccination60. Understanding how Mtb rewires the innate immune response at the early stages is crucial for the design of preventive and therapeutic vaccines, as well as novel therapies and disease biomarkers.

Revised version: Taken together, our data provide novel insights into the immunometabolic pathways involved in the trafficking of DCs to the lymph nodes, which may be targetable to generate sterilizing vaccine-induced immunity against TB. To this end, we propose that the development of tools aiming to improve the migratory capacity of DCs by fostering their HIF-1α-mediated glycolytic activity may increase the effectiveness of preventive strategies for TB, especially given that the number of DCs to reach the lymph node was shown to be a critical parameter for the outcome of DC-based vaccination61. Understanding how Mtb rewires the innate immune response at the early stages is crucial for the design of preventive and therapeutic vaccines, as well as novel therapies and disease biomarkers.

Methods

Chemical Reagents

LPS from Escherichia coli O111:B4 was obtained from Sigma-Aldrich (St. Louis, MO, USA). Dexamethasone (Dx) was from Sidus (Buenos Aires, Argentina). PX-478 2HCL was purchased from Selleck Chemicals (Houston, USA) and DMOG from Santa Cruz, Biotechnology (Palo Alto, CA, USA). Additionally, GSK2837808A was purchased from Cayman Chemical (Michigan, USA) together with echinomycin and sodium oxamate.

Bacterial strain and antigens

Earlier version: Mtb H37Rv strain was grown at 37°C in Middlebrook 7H9 medium supplemented with 10% albumin-dextrose-catalase (both from Becton Dickinson, New Jersey, USA) and 0.05% Tween-80 (Sigma-Aldrich). The Mtb γ-irradiated H37Rv strain (NR-49098) was obtained from BEI Resource (NIAID, NIH, USA).

Revised version: Mtb H37Rv strain was grown at 37°C in Middlebrook 7H9 medium supplemented with 10 % albumin-dextrose-catalase (both from Becton Dickinson, New Jersey, USA) and 0.05 % Tween-80 (Sigma-Aldrich). The Mtb γ-irradiated H37Rv strain (NR-49098) was obtained from BEI Resource (NIAID, NIH, USA). The RFP-expressing Mtb strain was gently provided by Dr. Fabiana Bigi (INTA, Castelar, Argentina).

Preparation of monocyte-derived DCs

Earlier version: Buffy coats from healthy donors were prepared at Centro Regional de Hemoterapia Garrahan (Buenos Aires, Argentina) according to institutional guidelines (resolution number CEIANM-664/07). Informed consent was obtained from each donor before blood collection. Monocytes were purified by centrifugation on a discontinuous Percoll gradient (Amersham, Little Chalfont, UK) as previously described31. Then, monocytes were allowed to adhere to 24-well plates at 5x105 cells/well for 1h at 37°C in warm RPMI-1640 medium (ThermoFisher Scientific, Waltham, MA). The mean purity of adherent monocytes was 85 % (range: 80-92%). The medium was then supplemented to a final concentration of 10% Fetal Bovine Serum (FBS, Sigma-Aldrich), human recombinant Granulocyte-Macrophage Colony-Stimulating Factor (10 ng/ml, GM-CSF, Peprotech, New Jersey, USA) and IL-4 (20 ng/ml, Biolegend, San Diego, USA). Cells were allowed to differentiate for 5-7 days (DC-SIGN+ cells in the culture > 90%).

Revised version: Buffy coats from healthy donors were prepared at Centro Regional de Hemoterapia Garrahan (Buenos Aires, Argentina) according to institutional guidelines (resolution number CEIANM-664/07). Informed consent was obtained from each donor before blood collection. Monocytes were purified by centrifugation on a discontinuous Percoll gradient (Amersham, Little Chalfont, UK) as previously described68. Then, monocytes were allowed to adhere to 24-well plates at 5×105 cells/well for 1 h at 37°C in warm RPMI-1640 medium (ThermoFisher Scientific, Waltham, MA). The mean purity of adherent monocytes was 85 % (range: 80-92 %). The medium was then supplemented to a final concentration of 10 % Fetal Bovine Serum (FBS, Sigma-Aldrich), human recombinant Granulocyte-Macrophage Colony-Stimulating Factor (10 ng/ml, GM-CSF, Peprotech, New Jersey, USA) and IL-4 (20 ng/ml, Biolegend, San Diego, USA). Cells were allowed to differentiate for 5-7 days (DC-SIGN+ cells in the culture > 90 %).

DC stimulation

Earlier version: DCs were stimulated with either irradiated Mtb (iMtb) or viable Mtb at equivalent OD600 doses for 24h at 37 °C. The cells were washed three times, and their phenotype and functionality were evaluated together with survival of activated cells; cell number and viability were determined by either trypan blue exclusion assays or MTT. Infections were performed in the biosafety level 3 (BSL-3) laboratory at the Unidad Operativa Centro de Contención Biológica (UOCCB), ANLIS-MALBRAN (Buenos Aires), according to the biosafety institutional guidelines.

Revised version: DCs were stimulated with either irradiated Mtb (iMtb) or viable Mtb at equivalent OD600 doses for 24 h at 37 °C. The cells were washed three times, and their phenotype and functionality were evaluated together with survival of activated cells; cell number and viability were determined by either trypan blue exclusion assays or MTT. Infections were performed in the biosafety level 3 (BSL-3) laboratory at the Unidad Operativa Centro de Contención Biológica (UOCCB), ANLIS-MALBRAN (Buenos Aires), according to the biosafety institutional guidelines.

DC treatments

Earlier version: When indicated, neutralizing monoclonal antibodies (mAb), or their corresponding isotype antibodies as mock controls, were added 30 min prior to DC stimulation to inhibit TLR2 (309717, Biolegend) or TLR4 (312813, Biolegend). In addition, DCs were incubated with PX-478 (20µM) or Echinomycin (1nM) with the purpose of inhibiting HIF-1α activity, DMOG (50µM) to stabilize HIF-1α, and oxamate (20mM) or GSK2837808A (20µM) to inhibit glycolysis. DC stimulation with iMtb occurred 30 min after treatment without drug washout.

Revised version: When indicated, neutralizing monoclonal antibodies (mAb), or their corresponding isotype antibodies as mock controls, were added 30 min prior to DC stimulation to inhibit TLR2 (309717, Biolegend) or TLR4 (312813, Biolegend). In addition, DCs were incubated with PX-478 (20 µM) or Echinomycin (1 nM) with the purpose of inhibiting HIF-1α activity, DMOG (50 µM) to stabilize HIF-1α, and oxamate (20 mM) or GSK2837808A (20 µM) to inhibit glycolysis. DC stimulation with iMtb occurred 30 min after treatment without drug washout.

Earlier version: In figure 6 and S6, dexamethasone-induced tolerogenic dendritic cells (Dx-DC) were generated by incubating DCs with 0.1µM of dexamethasone for 1h. Thereafter, cells were washed, and complete medium was added. Tolerogenic Dx-DCs were then stimulated or not with iMtb in the presence or not of DMOG (50µM).

Revised version: In figure 6 and S6, dexamethasone-induced tolerogenic dendritic cells (Dx-DC) were generated by incubating DCs with 0.1 µM of dexamethasone for 1 h. Thereafter, cells were washed, and “complete medium” was added. Tolerogenic Dx-DCs were then stimulated (or not) with iMtb in the presence or not of DMOG (50 µM).

Determination of metabolite concentrations

Earlier version: Lactate production and glucose concentrations in the culture medium was measured using the spectrophotometric assays Lactate Kit and Glicemia Enzimática AA Kit both from Wiener (Argentina), which are based on the oxidation of lactate or glucose, respectively, and the subsequent production of hydrogen peroxide32. The consumption of glucose was determined by assessing the reduction in glucose levels in culture supernatants in comparison with RPMI 10% FBS. The absorbance was read using a Biochrom Asys UVM 340 Microplate Reader microplate reader and software.

Revised version: Lactate production and glucose concentrations in the culture medium was measured using the spectrophotometric assays Lactate Kit and Glicemia Enzimática AA Kit both from Wiener (Argentina), which are based on the oxidation of lactate or glucose, respectively, and the subsequent production of hydrogen peroxide69. The consumption of glucose was determined by assessing the reduction in glucose levels in culture supernatants in comparison with RPMI 10 % FBS. The absorbance was read using a Biochrom Asys UVM 340 Microplate Reader microplate reader and software.

Quantitative RT-PCR

Earlier version: Total RNA was extracted with Trizol reagent (Thermo Fisher Scientific) and cDNA was reverse transcribed using the Moloney murine leukemia virus reverse transcriptase and random hexamer oligonucleotides for priming (Life Technologies, CA, USA). The expression of the genes HIF-1α and LDH-A was determined using PCR SYBR Green sequence detection system (Eurogentec, Seraing, Belgium) and the CFX Connect Real-Time PCR Detection System (Bio-Rad, CA, USA). Gene transcript numbers were standardized and adjusted relative to eukaryotic translation elongation factor 1 alpha 1 (EeF1A1) transcripts. Gene expression was quantified using the ΔΔCt method. Primers used for RT-PCR were as follows: EeF1A1 Fwd: 5’- TCGGGCAAGTCCACCACTAC -3’ and Rev: 5’-CCAAGACCCAGGCATACTTGA-3’; HIF-1α Fwd: 5’-ACTAGCCGAGGAAGAACTATGAA-3’ and Rev: 5’-TACCCACACTGAGGTTGGTTA-3’; and LDH-A Fwd: 5’-TGGGAGTTCACCCATTAAGC-3’ and Rev: 5’-AGCACTCTCAACCACCTGCT-3’.

Revised version: Total RNA was extracted with Trizol reagent (Thermo Fisher Scientific) and cDNA was reverse transcribed using the Moloney murine leukemia virus reverse transcriptase and random hexamer oligonucleotides for priming (Life Technologies, CA, USA). The expression of the genes HIF-1α and LDH-A was determined using the PCR SYBR Green sequence detection system (Eurogentec, Seraing, Belgium) and the CFX Connect Real-Time PCR Detection System (Bio-Rad, CA, USA). Gene transcript numbers were standardized and adjusted relative to eukaryotic translation elongation factor 1 alpha 1 (EeF1A1) transcripts. Gene expression was quantified using the ΔΔCt method. Primers used for RT-PCR were as follows: EeF1A1 Fwd: 5’-TCGGGCAAGTCCACCACTAC -3’ and Rev: 5’-CCAAGACCCAGGCATACTTGA-3’; HIF-1α Fwd: 5’-ACTAGCCGAGGAAGAACTATGAA-3’ and Rev: 5’-TACCCACACTGAGGTTGGTTA-3’; and LDH-A Fwd: 5’-TGGGAGTTCACCCATTAAGC-3’ and Rev: 5’-AGCACTCTCAACCACCTGCT-3’.

Immunofluorescence analysis

Earlier version: FITC-, PE- or PerCP.Cy5.5-labelled mAbs were used for phenotypic analysis of the following cell-surface receptor repertoires: FITC-anti-CD1a (clone HI149, eBioscience), PE-anti-DC-SIGN (clone 120507, RD System), PerCP.Cy5.5-anti-CD86 (clone 374216, Biolegend), FITC-anti-CD83 (clone HB15e, eBioscience), PE-anti-PD-L1 (clone MIH1, BD Pharmingen) and in parallel, with the corresponding isotype control antibody. Approximately 5x105 cells were seeded into tubes and washed once with PBS. Cells were stained for 30 min at 4°C and washed twice. Additionally, cells were stained for 40 min at 4°C with fluorophore-conjugated antibodies PE-anti-Glut1 (clone 202915 R&D Systems, Minnesota, USA) and in parallel, with the corresponding isotype control antibody. For HIF-1α determination, DCs were permeabilized with methanol and incubated with PE-anti-HIF-1α (clone 546-16, Biolegend). Stained populations were gated according to forward scatter (FSC) and side scatter (SSC) analyzed on FACScan (Becton Dickinson). Isotype matched controls were used to determine auto-fluorescence and non-specific staining. Analysis was performed using the FCS Express (De Novo Software) and results were expressed as median fluorescence intensity (MFI) or percentage of positive cells.

Revised version: FITC-, PE- or PerCP.Cy5.5-labelled mAbs were used for phenotypic analysis of the following cell-surface receptor repertoires: FITC-anti-CD1a (clone HI149, eBioscience), PE-anti-DC-SIGN (clone 120507, RD System), PerCP.Cy5.5-anti-CD86 (clone 374216, Biolegend), FITC-anti-CD83 (clone HB15e, eBioscience), PE-anti-PD-L1 (clone MIH1, BD Pharmingen) and in parallel, with the corresponding isotype control antibody. Approximately 5×105 cells were seeded into tubes and washed once with PBS. Cells were stained for 30 min at 4°C and washed twice. Additionally, cells were stained for 40 min at 4°C with fluorophore-conjugated antibodies PE-anti-Glut1 (clone 202915 R&D Systems, Minnesota, USA) and in parallel, with the corresponding isotype control antibody. For HIF-1α determination, DCs were permeabilized with methanol and incubated with PE-anti-HIF-1α (clone 546-16, Biolegend). Stained populations were gated according to forward scatter (FSC) and side scatter (SSC) analyzed on FACScan (Becton Dickinson). Isotype matched controls were used to determine auto-fluorescence and non-specific staining. Analysis was performed using the FCS Express (De Novo Software) and results were expressed as median fluorescence intensity (MFI) or percentage of positive cells.

Soluble cytokines determinations

Supernatants from DC populations or DC-T cell cocultures were harvested and assessment of TNF-α, IL-10, IL-17A or IFN-γ production was measured by ELISA, according to manufacturers’ instructions (eBioscience). The detection limit was 3 pg/ml for TNF-α and IL-17A, 6 pg/ml for IFN-γ, and 8 pg/ml for IL-10.

CD4+ T cell activation assay

Earlier version: Specific lymphocyte activation (recall) assays were carried out in cells from tuberculin purified protein derivate-positive skin test (PPD+) healthy donors by culturing DC populations and autologous T cells at a ratio of 10 T cells to 1 DC in round bottom 96-well culture plates for 5 days as detailed previously9. The numbers of DCs were adjusted to live cells before the start of the co-cultures. After 5 days, CD4+ T cell subsets were identified by immunolabeling according to the differential expression of CCR4, CXCR3, and CCR6 as previously reported33: CXCR3+CCR4−CCR6− (Th1), CXCR3−CCR4+CCR6− (Th2), CXCR3−CCR4+CCR6+ (Th17), and CXCR3+CCR4−CCR6+ (Th1* or Th1/Th17). The fluorochrome-conjugated antibodies used for flow cytometry analysis were CD4-FITC (clone A161A1, Biolegend), CXCR3-PE-Cy7 (clone G025H7, Biolegend), CCR4-PerCPCy5.5 (clone 1G1, BD Bioscience), CCR6-APC (clone 11A9, BD Bioscience), and CD3-APC-Cy7 (clone HIT3a, Biolegend). A viability dye, Zombie Violet (Biolegend), was used to exclude dead cells. Fluorescence Minus One (FMO) control was used to set proper gating for CXCR3-PE-Cy7, CCR4-PerCPCy5.5, and CCR6-APC detection. Cells were analyzed by fluorescence-activated cell sorting (FACS), using the BD FACSCANTO cytometer and FlowJoTM Software (BD Life Sciences).

Revised version: Specific lymphocyte activation (recall) assays were carried out in cells from tuberculin purified protein derivative-positive skin test (PPD+) healthy donors by culturing DC populations and autologous T cells at a ratio of 10 T cells to 1 DC in round bottom 96-well culture plates for 5 days as detailed previously9. The numbers of DCs were adjusted to live cells before the start of the co-cultures. After 5 days, CD4+ T cell subsets were identified by immunolabeling according to the differential expression of CCR4, CXCR3, and CCR6 as previously reported70: CXCR3+CCR4−CCR6− (Th1), CXCR3−CCR4+CCR6− (Th2), CXCR3−CCR4+CCR6+ (Th17), and CXCR3+CCR4−CCR6+ (Th1* or Th1/Th17). The fluorochrome-conjugated antibodies used for flow cytometry analysis were CD4-FITC (clone A161A1, Biolegend), CXCR3-PE-Cy7 (clone G025H7, Biolegend), CCR4-PerCPCy5.5 (clone 1G1, BD Bioscience), CCR6-APC (clone 11A9, BD Bioscience), and CD3-APC-Cy7 (clone HIT3a, Biolegend). A viability dye, Zombie Violet (Biolegend), was used to exclude dead cells. Fluorescence Minus One (FMO) control was used to set proper gating for CXCR3-PE-Cy7, CCR4-PerCPCy5.5, and CCR6-APC detection. Cells were analyzed by fluorescence-activated cell sorting (FACS), using the BD FACSCANTO cytometer and FlowJoTM Software (BD Life Sciences).

Chemotactic activity of DCs

Earlier version: Each DC population (4 x 105 cells in 75 µl) was placed on the upper chamber of a transwell insert (5 µm pore size, 96-well plate; Corning), and 230 µL of media (RPMI with 0.5 % FCS) with human recombinant CCL21 (200ng/ml) (Peprotech) was placed in the lower chamber. After 3 h, cells that had migrated to the lower chamber were removed and analyzed. The relative number of cells migrating was determined on a flow cytometer using Calibrite beads (BD Biosciences), where a fixed number of beads was included in each sample and the number of cells per 1,000 beads was evaluated. Data were normalized to the number of initial cells.

Revised version: Each DC population (4×105 cells in 75 µl) was placed on the upper chamber of a transwell insert (5 µm pore size, 96-well plate; Corning), and 230 µL of media (RPMI with 0.5 % FCS) with human recombinant CCL21 (200 ng/ml) (Peprotech) was placed in the lower chamber. After 3 h, cells that had migrated to the lower chamber were removed and analyzed. The relative number of cells migrating was determined on a flow cytometer using Calibrite beads (BD Biosciences), where a fixed number of beads was included in each sample and the number of cells per 1,000 beads was evaluated. Data were normalized to the number of initial cells.

In vivo migration assay

Earlier version: DCs were differentiated from bone marrow precursors obtained from naive BALB/c mice in the presence of murine GM-CSF (10ng/ml) and IL-4 (10ng/ml) both from Biolegend for 7 days. After differentiation, DCs were treated with oxamate (20mM) or PX-478 (10uM) and stimulated with iMtb. After 24h, DCs were stained with CFSE (5µM) and inoculated intradermically in the inguinal zone of naïve BALB/c mice. 3h post-injection, inguinal lymph nodes close to the site of inoculation were harvested and cells were stained with fluorophore-conjugated antibody PE-anti-CD11c (clone HL3, BD Pharmingen). Analysis was performed using the FlowJo Software and results were expressed as the percentage of CFSE+/CD11c+ cells.

Revised version: DCs were differentiated from bone marrow precursors obtained from naive BALB/c mice in the presence of murine GM-CSF (10 ng/ml) and IL-4 (10 ng/ml) both from Biolegend for 7 days. After differentiation, DCs were treated with oxamate (20 mM) or PX-478 (10 uM) and stimulated with iMtb. After 24 h, DCs were stained with CFSE (5 µM) and inoculated intradermally in the inguinal zone of naïve BALB/c mice. 3 h post-injection, inguinal lymph nodes close to the site of inoculation were harvested and cells were stained with fluorophore-conjugated antibody PE-anti-CD11c (clone HL3, BD Pharmingen). Analysis was performed using the FlowJo Software and results were expressed as the percentage of CFSE+/CD11c+ cells.

3D migration assay

Earlier version: 0.5x105 DCs were seeded on top of fibrillar collagen matrices polymerized from Nutragen 2mg/ml, 10% v/v MEM 10X (MEM invitrogen, Carlsbad, CA), UltraPure distilled water and 4-6% v/v bicarbonate buffer (pH=9) 7.5%. After 24h, cellular migration was quantified by taking images using an inverted microscope (Leica DMIRB, Leica Microsystems, Deerfield, IL) and the software Metamorph, as described previously34. Alternatively, in a similar manner, matrices were polymerized using Collagen (Sigma-Aldrich, C9791-10MG) in figure 5C. After 24h of cellular migration, matrices were fixed with PFA 4% during 30 min at room temperature and stained with DAPI (Cell signaling). Collagen was removed and membranes were mounted with DAKO. Images were taken by confocal microscopy (FluoView FV 1000), and cells were counted per field.

Revised version: 0.5×105 DCs were seeded on top of fibrillar collagen matrices polymerized from Nutragen 2 mg/ml, 10 % v/v MEM 10X (MEM invitrogen, Carlsbad, CA), UltraPure distilled water and 4-6 % v/v bicarbonate buffer (pH=9) 7.5 %. After 24 h, cellular migration was quantified by taking images using an inverted microscope (Leica DMIRB, Leica Microsystems, Deerfield, IL) and the software Metamorph, as described previously71. Alternatively, in a similar manner, matrices were polymerized using Collagen (Sigma-Aldrich, C9791-10MG) in figure 5C. After 24 h of cellular migration, matrices were fixed with paraformaldehyde (PFA) 4 % during 30 min at room temperature and stained with DAPI (Cell signaling). Collagen was removed and membranes were mounted with DAKO. Images were taken by confocal microscopy (FluoView FV 1000), and cells were counted per field.

Measurement of cell respiration with Seahorse flux analyzer

Earlier version: Bioenergetics were determined using a seahorse XFe24 analyzer. ATP production rates and relative contribution from the glycolysis and the OXPHOS were measured by the Seahorse XF Real-Time ATP Rate Assay kit. DCs (2 x 105 cells/well) were cultured in 3 wells per condition. The assay was performed in XF Assay Modified DMEM. Three consecutive measurements were performed under basal conditions and after the sequential addition of oligomycin and rotenone/antimycin (Agilent, USA). Extracellular acidification rate (ECAR) and oxygen consumption rate (OCR) were measured. Mitochondrial ATP production rate was determined by the decrease in the OCR after oligomycin addition. On the other hand, the complete inhibition of mitochondrial respiration with rotenone plus antimycin A, allows accounting for mitochondrial-associated acidification, and when combined with proton efflux rate (PER) data, allows calculation of Glycolysis ATP production rate. All OCR and ECAR values were normalized. Briefly, before the assay brightfield imaging was performed. Cellular area per condition was calculated by ImageJ software and imported into Wave (Agilent) using the normalization function.

Revised version: Bioenergetics were determined using a seahorse XFe24 analyzer. ATP production rates and relative contribution from the glycolysis and the OXPHOS were measured by the Seahorse XF Real-Time ATP Rate Assay kit. DCs (2×105 cells/well) were cultured in 3 wells per condition. The assay was performed in XF Assay Modified DMEM. Three consecutive measurements were performed under basal conditions and after the sequential addition of oligomycin and rotenone/antimycin (Agilent, USA). Extracellular acidification rate (ECAR) and oxygen consumption rate (OCR) were measured. Mitochondrial ATP production rate was determined by the decrease in the OCR after oligomycin addition. On the other hand, the complete inhibition of mitochondrial respiration with rotenone plus antimycin A, allows accounting for mitochondrial-associated acidification, and when combined with proton efflux rate (PER) data, allows calculation of Glycolysis ATP production rate. All OCR and ECAR values were normalized. Briefly, before the assay brightfield imaging was performed. Cellular area per condition was calculated by ImageJ software and imported into Wave (Agilent) using the normalization function.

SCENITH™ assay

Earlier version: SCENITH experiments were performed as previously described35 using the SCENITH kit containing all reagents and anti-puromycin antibodies (www.scenith.com). Briefly, DCs or PBMCs were treated for 40 min at 37°C in the presence of the indicated inhibitors of various metabolic pathways and puromycin. After the incubation, puromycin was stained using fluorescently labeled anti-Puromycin monoclonal antibody (clone R4743L-E8) with Alexa Fluor 647 or Alexa Fluor 488 and analyzed by flow cytometry. For metabolic analysis of monocyte subsets, PBMCs were labelled with PE-anti-CD16 (clone 3G8, Biolegend) and PECy7-anti-CD14 (clone HCD14, Biolegend) mAbs. The impact of the various metabolic inhibitors was quantitated as described35.

Revised version: SCENITH experiments were performed as previously described32 using the SCENITH kit containing all reagents and anti-puromycin antibodies (www.scenith.com). Briefly, DCs or PBMCs were treated for 40 min at 37°C in the presence of the indicated inhibitors of various metabolic pathways and puromycin. After the incubation, puromycin was stained using a fluorescently-labeled anti-Puromycin monoclonal antibody (clone R4743L-E8) with Alexa Fluor 647 or Alexa Fluor 488, and analyzed by flow cytometry. For metabolic analysis of monocyte subsets, PBMCs were labeled with PE-anti-CD16 (clone 3G8, Biolegend) and PECy7-anti-CD14 (clone HCD14, Biolegend) mAbs. The impact of the various metabolic inhibitors was quantitated as described32.

Transmission electron microscopy

Earlier version: DCs were fixed in 2.5% glutaraldehyde / 2% paraformaldehyde (EMS, Delta-Microscopies) dissolved in 0.1 M Sorensen buffer (pH 7.2) for 1h at room temperature, and then preserved in 1% paraformaldehyde (PFA) dissolved in Sorensen buffer. Adherent cells were treated for 1h with 1% aqueous uranyl acetate then dehydrated in a graded ethanol series and embedded in Epon. Sections were cut on a Leica Ultracut microtome and ultrathin sections were mounted on 200 mesh onto Formvar carbon coated copper grids. Finally, thin sections were stained with 1% uranyl acetate and lead citrate and examined with a transmission electron microscope (Jeol JEM-1400) at 80 kV. Images were acquired using a digital camera (Gatan Orius).

Revised version: DCs were fixed in 2.5 % glutaraldehyde / 2 % PFA (EMS, Delta-Microscopies) dissolved in 0.1 M Sorensen buffer (pH 7.2) for 1 h at room temperature, and then preserved in 1 % PFA dissolved in Sorensen buffer. Adherent cells were treated for 1 h with 1 % aqueous uranyl acetate then dehydrated in a graded ethanol series and embedded in Epon. Sections were cut on a Leica Ultracut microtome and ultrathin sections were mounted on 200 mesh onto Formvar carbon coated copper grids. Finally, thin sections were stained with 1 % uranyl acetate and lead citrate and examined with a transmission electron microscope (Jeol JEM-1400) at 80 kV. Images were acquired using a digital camera (Gatan Orius). For mitochondrial morphometric analysis, TEM images were quantified with the ImageJ “analyze particles” plugin in thresholded images, with size (μm2) settings from 0.001 to infinite. For quantification, 8–10 cells of random fields (1000x magnification) per condition were analyzed.

Changes of mitochondrial mass

Mitochondrial mass was determined in DCs by fixing the cells with PFA 4% and labeling them with the probe MitoSpy Green FM (Biolegend). Green fluorescence was analyzed by flow cytometry (FACScan, BD Biosciences).

Added in the revised version: GeneSet Enrichment Analysis of human monocytes

Added in the revised version: BubbleMap analysis was performed with 1,000 geneset-based permutations, and with “Signal2Noise” as a metric for ranking the genes. The results are displayed as a BubbleMap, where each bubble is a GeneSet Enrichment Analysis (GSEA) result and summarizes the information from the corresponding enrichment plot. The color of the Bubble corresponds to the population from the pairwise comparison in which the geneset is enriched. The bubble area is proportional to the GSEA normalized enrichment score (NES). The intensity of the color corresponds to the statistical significance of the enrichment, derived by computing the multiple testing-adjusted permutation-based p-value using the Benjamini–Yekutieli correction. Enrichments with a statistical significance above 0.30 are represented by empty circles.

Patient blood donors

Earlier version: TB patients were diagnosed at the División Tisioneumonología, Hospital F.J.Muñiz (Buenos Aires, Argentina) by the presence of recent clinical respiratory symptoms, abnormal chest radiography, and positive culture of sputum or positive sputum smear test for acid-fast-bacilli. Written, informed consent was obtained according to Ethics Committee from the Hospital Institutional Ethics Review Committee. Exclusion criteria included HIV positive patients and the presence of concurrent infectious diseases or comorbidities. Blood samples were collected during the first 15 days after commencement of treatment. All tuberculous patients had pulmonary TB (Table I).

Revised version: TB patients were diagnosed at the División Tisioneumonología, Hospital F.J.Muñiz (Buenos Aires, Argentina) by the presence of recent clinical respiratory symptoms, abnormal chest radiography, and positive culture of sputum or positive sputum smear test for acid-fast-bacilli. Written, informed consent was obtained according to the Ethics Committee from the Hospital Institutional Ethics Review Committee. Exclusion criteria included HIV positive patients and the presence of concurrent infectious diseases or comorbidities. Blood samples were collected during the first 15 days after commencement of treatment. All tuberculous patients had pulmonary TB (Table I).

Statistics

Earlier version: All values are presented as the median ± SEM of 3-13 independent experiments. Each independent experiment corresponds to 1 donor. For Seahorse assays, OCR and PER values are shown as mean ± SD. Comparisons between unpaired experimental conditions were made using either ANOVA for parametric data or Friedman test for non-parametric data followed by Dunn’s Multiple Comparison Test. Comparisons between paired experimental conditions were made using the two-tailed Wilcoxon Signed Rank Test for non-parametric data or T test for parametric data. A p-value of 0.05 was considered significant.

Revised version: All values are presented as the median ± SEM of 3-13 independent experiments. Each independent experiment corresponds to 1 donor. For Seahorse assays, OCR and PER values are shown as mean ± SD. Comparisons between unpaired experimental conditions were made using either ANOVA for parametric data or Friedman test for non-parametric data followed by Dunn’s Multiple Comparison Test. Comparisons between paired experimental conditions were made using the two-tailed Wilcoxon Signed Rank Test for non-parametric data or T test for parametric data. Correlation analyses were determined using the Spearman’s rank test. A p-value of 0.05 was considered significant.

Study approval

Human specimens

The study design was reviewed and approved by the Ethics Committees of the Academia Nacional de Medicina (49/20/CEIANM) and the Muñiz Hospital, Buenos Aires, Argentina (NI #1346/21). All participants voluntarily enrolled in the study by signing an informed consent form after receiving detailed information about the research study.

Mouse studies

All experimental protocols were approved by the Institutional Animal Care and Use of the Experimentation Animals Committee (CICUAL number 090/2021) of the Institute of Experimental Medicine (IMEX, Buenos Aires).

Supporting information